Cellular modifiers of TDP-43 phase transition and cytoplasmic aggregation

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

This revised manuscript represents a partial response to the concerns raised in the first round of review. The authors have made one genuine mechanistic addition in the form of the semi-permeabilized cell reconstitution assay, removed the most overreaching conclusions regarding the contribution of cytoplasmic TDP-43 aggregation to disease, and made several minor presentational improvements. However, the central weaknesses of the original submission remain substantially unaddressed. The exclusive reliance on non-physiological TDP-43 variants, the incompletely resolved mechanism linking XPO1 to TDP-43 phase behavior, and the limited organoid validation continue to limit confidence in the major claims. The authors have, in several instances, responded by removing contested data rather than by providing the additional evidence that was requested.

(1) The justification for the 2KQ acetylation-mimetic system remains inadequate.

The authors respond to the concern about the non-physiological nature of the 2KQ mutant by citing published evidence that TDP-43 acetylation occurs in ALS patient spinal cord and is upregulated under oxidative and proteotoxic stress conditions. While these references are real and support the relevance of acetylation as a pathological post-translational modification, they do not resolve the central concern: there is no quantification of how much endogenous TDP-43 is acetylated at the specific lysine residues mimicked by 2KQ in degenerating human neurons, and no evidence that the degree of RNA-binding disruption imposed by the double glutamine substitution is ever achieved by endogenous acetylation in vivo. The 2KQ mutant eliminates RNA binding essentially completely, whereas physiological acetylation events are graded, reversible, and likely partial. The response conflates the existence of TDP-43 acetylation as a phenomenon with validation that 2KQ is a physiologically accurate model of that phenomenon. None of the new experiments address the request to test whether wild-type TDP-43 expressed at near-physiological levels, or a bona fide heterozygous ALS-linked TARDBP mutant in iPSC-derived neurons, responds to XPO1 modulation in a qualitatively similar fashion. Until this is shown, the mechanistic conclusions of this paper remain constrained to a highly artificial overexpression system and cannot be extrapolated to physiological or pathological TDP-43 biology with confidence.

We agree with the reviewer that the TDP-43 2KQ mutant is a non-physiological variant. To address this concern, we have removed all statements that extrapolate our findings to disease pathogenesis. As reflected in the revised title, we now present this study as an investigation of factors that modulate TDP-43 phase transition and aggregation using a sensitized model system, rather than as a direct disease model. The Abstract and Discussion have also been revised accordingly.

The choice of the 2KQ mutant was dictated by the requirements of the screening strategy. To identify modulators of TDP-43 phase transition, it was necessary to use a TDP-43 variant that reliably undergoes phase separation in cells within an experimentally practical timeframe. In this context, we believe the 2KQ mutant provides a suitable and justified experimental tool. We have revised the manuscript to clearly distinguish observations made with this engineered construct from conclusions regarding physiological or disease-associated TDP-43.

(2) The homozygous K181E organoid model is still not adequately justified, and no heterozygous comparison has been provided.

The authors acknowledge that the homozygous background is "more sensitive for detecting phospho-TDP-43" and argue that homozygous conditions are commonly used in experimental TDP-43 research. However, the critical issue is not whether homozygous models are used in general, but whether the homozygous background specifically alters the relative contribution of cytoplasmic aggregation versus nuclear RNA-processing dysfunction in this study. In a homozygous K181E model, both alleles produce an RNA-binding-defective TDP-43, meaning that every molecule of endogenous TDP-43 in the cell is dysfunctional. This is categorically different from the patient situation in which one wild-type allele is present, and it may substantially exaggerate nuclear loss-of-function relative to cytoplasmic gain-of-function phenotypes. The authors have not performed the requested comparison with heterozygous K181E/+ organoids, nor have they acknowledged that the organoid genotype itself could bias the interpretation of what KPT-276 treatment rescues. Given that the organoid section is now the sole in-disease-model validation of the XPO1 mechanism, this limitation is more consequential than it was in the original submission.

We agree with the reviewer and have removed all speculative statements regarding the relative contributions of cytoplasmic aggregation and RNA splicing defects to disease pathogenesis. The organoid section has also been revised to focus solely on the experimental findings. Specifically, we only present evidence that inhibition of nuclear export in a sensitized organoid model promotes the accumulation of cytoplasmic phosphorylated TDP-43 without making broader claims regarding its role in disease pathogenesis.

(3) The new semi-permeabilized cell data is a genuine contribution, but the mechanistic interpretation remains insufficiently constrained.

The development of the streptolysin O semi-permeabilized cell reconstitution system is the most substantive new addition to this revision. The finding that LMB-stabilized anisosomes resist cytosol washout but dissolve upon RNase T1 treatment is interesting and provides a plausible indirect mechanism: XPO1 inhibition retains nuclear RNA, and this elevated nuclear RNA availability contributes to maintaining the liquid LLPS state of the TDP-43 2KQ condensate. This is a meaningful mechanistic advance and deserves credit. However, several important limitations of this new data are not adequately discussed. First, RNase T1 degrades single-stranded RNA globally during permeabilization, so the experiment does not identify which specific RNA species stabilize the anisosome, nor whether these are pre-mRNA splicing intermediates, mature mRNA, non-coding RNA, or another class. Second, the same nuclear export blockade that retains RNA will also retain the nuclear concentrations of many RNA-binding proteins, splicing factors, and other XPO1-dependent cargos. The RNase T1 experiment does not exclude the possibility that the relevant effect is mediated by an RNA-binding protein whose nuclear concentration increases upon LMB treatment and which, upon RNase digestion, can no longer engage TDP-43 or the anisosome shell. Third, the permeabilized cell system is by definition not intact and has lost cytosolic factors; whether the RNA-dependent stabilization of anisosomes operates in the same way in intact cells during physiological or pathological nuclear export perturbation is an assumption, not a demonstrated fact. The authors should more carefully frame these data as hypothesis-generating and explicitly note these alternative interpretations in the Discussion.

We have now added some sentences on page 11 to acknowledge the limitation of our experiments. It reads as “However, our data does not exclude other RNA species or RNA-binding proteins as anisosome stabilizer. Whether RNA-dependent stabilization of anisosomes operates in the same way in intact cells also requires further validation.”

(4) The conceptual asymmetry between XPO1 inhibition and XPO1 overexpression phenotypes is not resolved by the new mechanism.

The paper continues to present two XPO1 perturbation phenotypes that are difficult to reconcile within a single mechanistic model. XPO1 inhibition enlarges anisosomes, maintains their liquid character by FRAP, and retains them in the nucleus. XPO1 overexpression also enlarges TDP-43 puncta, but these are FRAP-impaired, gel-like, and appear in the cytoplasm. The RNA-retention model proposed by the new semi-permeabilized data explains why XPO1 inhibition stabilizes the liquid state, but it does not explain why XPO1 overexpression drives the opposite outcome: gel-like hardening and cytoplasmic redistribution. If increased nuclear RNA availability is the key variable downstream of XPO1 inhibition, then XPO1 overexpression would be expected to decrease nuclear RNA and thereby destabilize anisosomes toward dissolution or hardening. The paper does not test whether nuclear RNA levels are indeed altered by XPO1 overexpression, nor whether the cytoplasmic gel-like puncta seen in XPO1-overexpressing cells are RNA-poor relative to control anisosomes. The revised Discussion does not engage with this asymmetry in a satisfying way, and the figure model remains qualitative. A quantitative or at least semi-quantitative model that accounts for both arms of the XPO1 perturbation is needed.

We thank the reviewer for this point. We have now explicitly mentioned in the discussion that the effect of XPO-1 on anisosome dynamics is likely mediated by an indirect mechanism. We also acknowledge that we do not fully understand why over-expression of XPO-1 causes TDP-43 to accumulate in gel-like structures in the cytoplasm. Although we did not check whether overexpression of XPO1 increases cargo export, we cited studies showing that over-expressed XPO1 disrupts the normal distribution of cargos between nucleus and cytoplasm on page 7. To avoid confusion, we also revised the result part on page 7, emphasizing on the difference rather than the similar increase in puncta size by opposing manipulations.

(5) The removal of RNA-seq data weakens rather than strengthens the organoid section.

The authors have removed the bulk RNA-seq analysis from the revised manuscript in response to concerns that the modest transcriptional rescue was being over-interpreted. While the decision to remove over-interpretation is appropriate, the result is that the organoid section now rests entirely on pTDP-43 immunostaining as its sole readout. The revised paper thus uses reduction in immunofluorescent pTDP-43 puncta in homozygous K181E organoids as the only evidence that nuclear export inhibition mitigates TDP-43 proteinopathy in a disease-relevant context. This is a weaker evidentiary base than before the revision, not an improvement. The originally requested more sensitive orthogonal readouts, including biochemical fractionation for SDS-insoluble TDP-43, filter-trap assays, or RNA aptamer-based detection of TDP-43 aggregates, remain absent. Without at least one additional independent measure confirming that cytoplasmic TDP-43 aggregation is genuinely reduced rather than simply rendered antigenically undetectable, the organoid conclusion is not adequately supported. At minimum, the authors should provide total and cytoplasmic TDP-43 fractionation data from organoid lysates to corroborate the immunostaining result.

We removed the RNA-seq analysis because both the reviewers and editors agreed that the modest transcriptional rescue should not be overinterpreted. Upon reconsideration, we believe that reinstating these data would not address the reviewer's principal concern, namely whether nuclear export inhibition reduces TDP-43 aggregation in organoids. We have therefore chosen to limit our conclusions to the direct observation supported by the current data, namely a reduction in cytoplasmic phosphorylated TDP-43 immunoreactivity. We also add a sentence to acknowledge that “Whether the reduction in pTDP-43 immunoreactivity reflects a decrease in insoluble TDP-43 aggregates remains to be determined” on page 10.

(6) No functional neuronal readout has been provided for the organoid model.

The organoid section now makes the claim that "nuclear export is required for the formation of p-TDP-43-containing aggregates in a disease-relevant organoid model," but no measure of neuronal health, integrity, or function is reported in association with this. Even a simple assessment of neuron survival by TUJ1 or MAP2 quantification, neurite complexity, or cleaved caspase-3 staining before and after KPT-276 treatment would substantially strengthen the biological significance of the pTDP-43 reduction. The current data establish a pharmacological effect on a pathological marker but do not demonstrate that this has any consequence for neuronal biology in the organoid, which is what the disease-relevance framing implies.

We thank the reviewer for this helpful suggestion. We agree that assessments of neuronal survival or function would be important if the manuscript were claiming that nuclear export inhibition improves neuronal health or rescues disease phenotypes in the organoid model. However, in response to the reviewers' comments regarding the physiological relevance of the homozygous K181E organoids, we have substantially revised both the framing and interpretation of this section.

Specifically, we have revised the statement to read, "These results imply that maintaining TDP-43 in the nuclear demixed liquid state might diminish p-TDP-43 accumulation but whether the reduction of pTDP-43 immunoreactivity reflects a decrease in insoluble TDP-43 aggregates remains to be determined," thereby limiting our conclusion to the direct experimental observation. We no longer make claims regarding disease modification or functional rescue in the organoid model. Given this revised scope, we believe that additional measurements of neuronal survival or function, while certainly of interest, are not essential to support the conclusions presented in this study. We have also revised the conclusion to explicitly acknowledge that the functional consequences of reducing p-TDP-43-positive puncta (whether this can be translated to reduced aggregation) remain to be determined by future studies (page 10).

(7) The abstract and title continue to overstate the mechanistic conclusions.

Despite the stated intent to reframe the study as a screening study and to temper the conclusions, the revised abstract retains the language: "These findings establish nuclear export as a key regulator of TDP-43 phase transitions and define a mechanistic framework that links altered nuclear transport and phase dynamics to TDP-43 aggregation potential." Similarly, the Discussion still states: "a particularly compelling aspect of our study is the discovery that the nuclear export receptor XPO1 governs TDP-43 liquid-to-solid transitions and subcellular localization." The word "governs" and the phrase "establish nuclear export as a key regulator" are not warranted by data that derive entirely from an overexpressed acetylation-mimetic mutant in a colon cancer cell line and a homozygous K181E organoid model. A more accurate framing would describe these findings as identifying nuclear export as one of several cellular processes that modulate TDP-43 phase behavior in a sensitized model system, with an indirect RNA-mediated mechanism that remains to be defined at the molecular level. The title change from "governs" to "modulates" is appreciated but does not extend into the abstract and Discussion, where the strong causal language persists.

We have revised the title of the paper, reframing it as a screen that reveals modulators of TDP-43 phase separation. The last sentence of the abstract is also revised accordingly. It now reads as “These findings identify multiple modulators of TDP-43 phase transitions in a sensitized model system and establish a framework for further dissecting the link between nuclear transport and TDP-43 phase dynamics.” We also tone down our conclusions and discussions.

(8) Individual siRNA knockdown validation for XPO1 has not been provided.

The authors argue that validation with 6 independent siRNAs across two rounds of screening, combined with convergent pharmacological data, is sufficient to establish XPO1 as a genuine hit. While the convergence of chemical and genetic evidence is reassuring, the specific request was for protein-level confirmation of XPO1 knockdown efficiency in the DLD1 TDP-43 2KQ cells used for mechanistic follow-up, together with demonstration that the anisosome phenotype is specifically caused by loss of XPO1 and not by off-target effects. This is a straightforward experiment, and its absence is particularly notable given that the entire mechanistic XPO1 narrative hinges on this specificity. At minimum, an immunoblot confirming XPO1 protein depletion in cells treated with the siRNA pool identified in the screen, in the same cell background and induction conditions as the follow-up experiments, should be provided.

While we agree with the reviewer that studies relying on siRNA should provide sufficient information regarding knockdown efficiency and specificity, we respectfully disagree that this should be a major concern in the present study. As explained in the manuscript, we deliberately chose not to pursue mechanistic studies using chronic XPO1 knockdown because prolonged depletion of this essential nuclear export factor is likely to produce secondary effects that could complicate data interpretation. Instead, we employed multiple chemically distinct XPO1 inhibitors to achieve acute inhibition, thereby minimizing indirect consequences while providing a more appropriate approach for mechanistic analysis.

We agree that assessing knockdown efficiency is technically straightforward. However, because our mechanistic conclusions are based primarily on acute pharmacological inhibition rather than siRNA-mediated depletion, we prioritized experiments that directly addressed the central mechanistic questions raised by the reviewers, particularly the semi-permeabilized cell assay. Moreover, the XPO1 inhibitors used in this study are well-characterized, highly specific compounds that have been extensively validated and widely used in the literature. We therefore believe that our experimental strategy provides a reliable basis for the conclusions presented.

(9) The identity of XPO1-dependent cargos that regulate anisosome dynamics remains entirely unknown.

The authors acknowledge that XPO1 does not directly bind TDP-43 and that the mechanism is likely indirect. The new RNA data provides one plausible indirect pathway. However, the possibility that one or more specific RNA-binding proteins or splicing factors, whose nuclear levels rise upon XPO1 inhibition, are the proximate drivers of anisosome stabilization has not been addressed. This matters because if the relevant mechanism operates through a specific cargo rather than bulk RNA retention, the model for how nuclear export connects to TDP-43 aggregation in disease would be fundamentally different. The authors decline to pursue adaptor identification on grounds of scope, which is a defensible position for future work. However, the framing should explicitly state that the current data cannot distinguish between bulk RNA retention and cargo-specific effects, and that the conclusion that nuclear export modulates TDP-43 phase behavior via RNA accumulation is a working hypothesis supported by but not proven by the RNase T1 experiment.

We thank the reviewer for this helpful suggestion. We have now added a sentence on page 11, which state that “our data does not exclude other RNA species or RNA-binding proteins as anisosome stabilizer. Whether RNA-dependent stabilization of anisosomes operates in the same way in intact cells also requires validation.”

Minor remaining issues.

The number of independent iPSC clones and organoid batches used for the KPT-276 treatment experiment is now stated as two batches per condition, which is minimal for a 3D organoid study and does not fully address the concern about clone-level variability. Ideally, organoids from at least two independently derived isogenic clones per genotype would be used. The mCherry overexpression control added in Supplemental Figure 4 is a useful addition and is acknowledged. The immunoblotting confirmation that drug treatments do not alter total TDP-43 levels addresses a prior concern adequately. The addition of the sentence noting that anisosomes have not been validated in human patient samples is appreciated and appropriate. Statistical detail has been improved in figure legends. These minor improvements are noted positively but do not compensate for the major unresolved concerns above.

We thank the reviewer for his/her appreciation of our previous revision. We hope that the new changes now satisfactorily address the remaining concerns.

Reviewer #2 (Public review):

This manuscript addresses an important and timely question in TDP-43 biology by systematically identifying regulators of TDP-43 anisosome formation, with a particular focus on nuclear export via XPO1. Using a combination of unbiased chemical screening, genetic perturbation, and advanced imaging approaches, the authors propose that inhibition of nuclear export modulates the abundance and biophysical properties of TDP-43 anisosomes. They further strengthen their findings by introducing an additional model system, a semi-permeabilized in vitro assay, which provides mechanistic evidence that XPO1 activity prevents anisosome dissolution by retaining nuclear RNAs. The study is conceptually innovative and has potential relevance for neurodegenerative diseases characterized by TDP-43 pathology. Some minor concerns remain, mostly about experimental design of the newly added data.

Strengths:

(1) The study employs an unbiased, hypothesis-free compound screen to identify regulators of TDP-43 anisosome formation, which is a major strength and reduces confirmation bias.

(2) The authors combine chemical and genetic screening approaches, providing orthogonal validation of key pathways and increasing confidence in the biological relevance of top hits.

(3) The focus on biophysical properties of TDP-43 assemblies, assessed through imaging and FRAP, moves beyond simple presence/absence of aggregates and provides mechanistic insight into the biophysical states of TDP-43.

(4) The use of multiple experimental modalities, including live-cell imaging, FRAP, pharmacological perturbation, and transcriptomic analysis, reflects a technically sophisticated and ambitious study design.

(5) The authors attempt to extend findings beyond immortalized cancer cell lines by incorporating organoid models, demonstrating awareness of disease relevance and translational importance.

(6) The authors extend their study by incorporating a semi-permeabilized in vitro system, which provides compelling evidence that inhibition of nuclear export promotes the retention of nuclear anisosomes, an effect driven by the accumulation of nuclear RNAs.

Overall, the manuscript is clearly written and logically structured, making complex experimental workflows accessible and the central hypotheses easy to follow.

We thank the reviewer for acknowledging the strength and the potential significance of our study.

Weaknesses:

(1) The manuscript has significantly improved with the revisions. Some experimental procedures and method details, as well has statements remain incompletely described:

(a) What is the smear in Figure S1 after VLX treatment?

We thank the reviewers for the positive assessment. We do not know why VLX treatment causes a fraction of TDP-43 to migrate slowly. We suspect that it may form detergent-insoluble aggregates. However, we cannot be sure whether this occurred during drug treatment or sample preparation. We now add a sentence in the figure legend to clarify this point.

(b) The authors state that "The reduction in TDP-43 signal was not due to protein elimination.", however no data is provided to support that statement.

We reasoned that the reduction in TDP-43 was probably not caused by protein elimination because the puncta could be reformed when permeabilized cells were incubated with exogenously added cytosol and ATP/GTP. We have revised the text to avoid this confusion. The revision on page 8 reads as “The reduction in TDP-43 signal probably resulted from a shift of TDP-43 from a phase-separated high fluorescent state into a soluble state with reduced fluorescence intensity (Zhang et al., 2026). We attributed this phenotype to the depletion of cytosolic factors and ATP during cell permeabilization because it is known that anisosome formation and maintenance require HSP70, a cytosolic ATPase (Yu et al., 2021).”.

(c) The authors state that "TDP-43 shifts from phase-separated state to a soluble state ...", however no data is provided to support that statement.

Since TDP-43 protein was apparently still in the nucleus after cell permeabilization (see above) but became invisible, the best interpretation is that the protein is shifted into a soluble state, which reduces the fluorescence intensity substantially. We have revised the text to clarify this point. We also cited a recent study showing that EGFP-alpha-synuclein oligomerization/aggregation enhances its fluorescence intensity in cells.

(d) Why did the authors choose cow lover cytosol for this study?

The main reason is because we have access to a large amount of cow liver cytosol that is known to have activities in in vitro reconstitution assays. We now cite several papers from us that reported the use of the same cytosol in other in vitro assays in the method (page 14).

(e) The experimental setup for supplementing with cytosol/ATP/GTP is unclear. A more detailed schematic would be helpful to understand at what stage in the experiment these factors were added. Which step of the protocol was performed at 37 {degree sign}C, which is indicated in the figure schematic but not described in the methods.

We now revise the schematic in Figure 6A and include more details in the method and figure legend.

(f) In the organoid model, the authors mention that they observe similar levels of total TDP-43, however they do not provide quantification. Instead, they provide a graph that shows highly significant changes in nuclear TDP-43, which was not addressed in the text.

The total TDP-43 level was shown by immunostaining in green in Figure 7. This was used as a control to show that the increase in p-TDP-43 was not simply caused by an overall increase in its protein level. We have added the quantification to Figure 7C. We also discuss the reduced nuclear TDP-43 in organoids bearing the disease mutation in the main text.

Additionally, some questions remain unclear:

(1) The anisosomes induced by ATP/GTP or cytosol are insufficiently characterized. It remains unclear whether these structures correspond to canonical ring-shaped anisosomes, and whether they exhibit dynamic (liquid-like) or more static (gel-like) properties.

We agree that the structures reformed after incubating permeabilized cells with cytosol and ATP/GTP are not fully characterized. Due to their small size, we could not see the typical ring-shaped anisosome morphology. FRAP experiment is also tricky. Due to these issues, we have revised the text to acknowledge that we do not know the exact identity of these structures. We speculate that they are anisosome-related because like anisosome formation, it depends on cytosolic factor and energy (page 9). It is worth noting that whether these structures are anisosomes is not the main conclusion of this experiment. We conclude from this experiment that TDP-43 was still in the nucleus after cell permeabilization (not degraded). The fact that we could not see the protein likely because the protein was in a low-fluorescence soluble state.

(2) The contribution of the cytosol and ATP/GTP supplementation experiments to the overall narrative is unclear. While the findings are intriguing, their interpretation within the context of the study is not well articulated. In particular, the rationale for including cytosol is not sufficiently justified, given that ATP/GTP alone induces a pronounced effect, whereas cytosol alone does not.

Since the formation of anisosome requires HSP70, a cytosolic chaperone that likely needs to be imported into the nucleus, we included cytosol and ARS/GTP in our in vitro reaction. We revise the description in the result part to improve clarity (page 8-9).

(3) The authors should address why endogenous XPO1 does not co-localize with anisosomes, whereas overexpressed XPO1 does. This raises the possibility that the observed co-localization may be an artifact of non-physiological protein levels, which should be discussed.

As discussed in Yu H et al., Science 2021, proteins in anisosomes cannot be stained by antibodies due to an antibody accessibility issue. It was mentioned in our paper as “since antibody staining could not conclusively demonstrate the sequestration of endogenous XPO-1 in anisosomes due to an antibody penetration barrier {Yu, 2021 #937}.” We now revise this section completely to better clarify this point. We could see overexpressed XPO1 in anisosome because it has a mCherry tag.

(4) The iPSC-based model remains insufficiently characterized. While the authors propose that this system recapitulates the accumulation of liquid and solid aggregates resembling anisosomes, it is unclear whether this phenotype is robustly observed and whether KPT treatment effectively modulates it.

The full characterization of the iPSC-derived organoids is presented in a second paper that is posted in BioRxiv (https://www.biorxiv.org/content/10.1101/2025.11.09.687455v2), which is cited. This manuscript reports not only the accumulation of p-TDP43, but also other ALS-related phenotypes including cell death, gene transcriptional changes, cryptic exon inclusion etc. in mutant organoids.

(5) The rationale for the selected treatment durations is unclear, and the timing appears inconsistent across experiments (ranging from 3 to 16 hours), including within experiments involving the same compound. This variability should be justified or standardized.

The longer treatment (24 h) was used in the chemical genetic screen in which different drugs may act with different efficiency. To maximize our chance of detecting more drug effect, we used a longer treatment scheme. For later follow-up experiments involving Spuatin-1, Bortezomib, TRP, because the phenotype appears quickly. To avoid secondary effects from long treatment, we shortened the treatment to 3-5 hours. For LMB treatment, we used long treatment to reveal the steady-state phenotype (anisosome enlargement in size and reduction in number has reached maximum). This time point was determined in Figure 5A-C. In contrast, shorter treatment (5 h) was to reveal early changes that might be causal to the end-point phenotypes (e.g. the anisosome fusion phenotype could be detected as early as 5 h post-treatment). We have added some explanations in the result section to make this point clear.

(6) Several figure legends require clarification:

We thank the reviewer for pointing out the inconsistencies. We have corrected the outstanding issues, as explained below.

(a) In the section stating “Collectively, our results suggest that the stability and dynamics of anisosomes are modulated by XPO1-mediated nuclear export ...", the cited figure appears to be incorrect. This should refer to Figure 5L rather than Figure 5J.

Thanks for pointing out this error. This is now corrected.

(b) Figure 1B: Please specify the number of replicates per concentration, the number of cells analyzed, and the model used for regression analysis. Additionally, the legend indicates a treatment duration of 15 hours, whereas Figure 1A states 24 hours.

Due to the large sample size, each concentration was analyzed once. We have added other information to the figure legend. We also remove the redundant inaccurate information from the figure legend. The treatment time shown in the figure is correct as it was also indicated in the method.

(c) Figure 2G: The authors state "7 anisosomes per condition," but the graph displays only 4-6 data points. Please clarify what each data point represents.

We thank the reviewer for noticing the discrepancy and apologize for the error. We have corrected the figure legend to indicate that 4-6 anisosomes were analyzed for each condition. In Figure 2G, each data point represents the initial fluorescence loss rate averaged from the first 10 sec after reverse photobleaching.

(d) Figures 3B and 3G: Please clarify whether a defined threshold was used to determine a "reduction in anisosome number."

In Figure 3B, we used Z score >2 as the threshold. This is now mentioned in the legend and defined in the method. There is no Figure 3G.

(e) Figure 4B: These do not represent biological replicates, as all samples derive from a single cell line; rather, they constitute independent experimental replicates.

We have changed the figure legend throughout the paper accordingly.

(f) Figures 5B and 5H: The legend states "n = 3 biological repeats," but the number of data points shown appears higher. Please clarify.

In Figure 5B, the graph reflects data collected from 3 independent replicates. To ensure reliable baseline measurement, for each experiment, two independent control samples were included, which is why it has 6 data points. In Figure 5H, each dot represents a randomly selected imaging field. We now mention the total number of fields analyzed.

(g) Figures 5K, 6C, and 6E: "Mean Fluorescence Intensity (MPI)" should be corrected to "MFI."

These are all fixed. Thank you for pointing this out.

(h) Figure 6C: Please include the number of cells analyzed and provide relevant statistical measures (e.g., R2, p-value).

We now include the cell number in the legend and R2 and p-value in the figure.

(i) Figure 6D: The experimental timeline is unclear. Please specify the duration of incubation and the timing of each step.

We now revise the experimental scheme in Figure 6A to better explain the experiment and the sequence of different events. For Figure 6D, permeabilized cells were incubated with cytosol with or without ARS/GTP for 40 min. This information is added to the figure legend.

(j) Figure 7B: Improved labeling is needed (e.g., clarification of "mean spot volume") to better align with the figure legend.

To improve clarity, we change mean spot volume to p-TDP-43 puncta mean volume. This refers to the average volume of segmented phosphorylated TDP-43-positive puncta.

Reviewer #3 (Public review):

Summary:

TDP-43 proteinopathy is broadly found in neurodegenerative diseases. This manuscript investigates how nuclear export influences the biophysical properties of TDP-43. The authors use a combination of chemical screening and genome-wide siRNA screening to identify pathways that modulate TDP-43 liquid-to-solid transitions. Overall, the study employs a broad array of approaches and addresses an important question in TDP-43 pathobiology. The identification of nuclear export as a central regulator is compelling and conceptually aligns with the emerging view that TDP-43 nucleocytoplasmic trafficking is a major defect in neurodegeneration.

Strengths:

This work integrates chemical and genetic screening to identify novel modifiers. The candidates were validated in both reporter cell lines and iPS-differentiated organoids. The findings support the nucleocytoplasmic transport is important for the biophysical properties of TDP-43.

Comments on revised version.

The manuscript has been improved with more data and clarification. The RNase T1 treatment experiment suggests that RNA is required for anisosome integrity. However, this does not directly demonstrate LMB increases nuclear RNA availability as changes in protein composition or other RNA-dependent mechanisms may also contribute. The conclusion and discussion need to be edited to consider these alternative scenarios. Overall, as most of the evidence remains indirect, the manuscript should avoid overinterpretation regarding the mechanisms underlying TDP-43 phase transition and aggregation.

We thank the reviewer for this helpful suggestion. We have added a few sentences in the discussion (page 10) to acknowledge the limitation of the semi-permeabilized cell assay. Specifically, we mentioned that “However, our data does not exclude other RNA species or RNA-binding proteins as anisosome stabilizer. Whether RNA-dependent stabilization of anisosomes operates in the same way in intact cells also requires further validation.” We also revise our manuscript throughout to avoid over-interpretation.

Recommendations for the authors:

Editor's notes:

The value of the work is clear.

We also recognise that it may not be possible/practical to get around the 'incomplete' appellation attached to this body of work, by further experiments. However, there may be scope here to retreat from the less well supported mechanistic claims-by editing the title, abstract and discussion and thus earn a 'solid' descriptor on a revised paper that remains a useful addition.

The authors are best placed to consider creatively how to achieve this.

We thank the editors for this helpful suggestion. We have revised the manuscript extensively to address every single concerns of the reviewers.

Reviewer #1 (Recommendations for the authors):

This revision addresses some minor concerns and adds one mechanistic experiment of genuine value. However, the major deficiencies of the original submission persist: the exclusive reliance on non-physiological TDP-43 model systems without validation in more disease-relevant contexts, the unresolved asymmetry between the two XPO1 perturbation phenotypes, the thin organoid section that now has fewer readouts than before the revision, and overstatement of the mechanistic conclusions in the abstract and Discussion. The manuscript in its current form still does not provide methods, data, and analyses that sufficiently support the primary claim that nuclear export is an established key regulator of TDP-43 phase transitions with mechanistic and disease relevance.

As mentioned before, we have clarified the interpretation of the data, tempered conclusions where appropriate, and revised the text to explicitly acknowledge the limitations of the current study. We hope that these changes satisfactorily addressed the reviewer’s concern.

Reviewer #2 (Recommendations for the authors):

I would suggest adjusting the title to match the data, which shows so much more than just an effect of nuclear export.

We thank the reviewer for this suggestion. We have changed the title to “Cellular modifiers of TDP-43 phase transition and cytoplasmic aggregation”

When introducing the semi-permeabilized cell-based in vitro assay, it would be helpful to add a short statement describing what this model resembles, and what advantage it can bring to use this system in the context of the study.

We have revised this section extensively and hope that improves the clarity. See marked text in page 8-9.

Changes from the earlier version to the revised version

Abstract

Earlier version: RNA-binding protein TAR DNA-binding protein 43 (TDP-43) can form liquid-like, nuclear assemblies whose phase behavior may influence its aggregation propensity and neurotoxic activity. The mechanism(s) that modulates the transition of TDP-43 from a liquid to solid phase is poorly defined. Here we combine chemical and genome-wide genetic screenings to identify cellular factors that modulate the phase behavior of an RNA-binding defective TDP-43 mutant that mimics an Amyotrophic Lateral Sclerosis (ALS)-associated variant. Our screens uncover multiple cellular processes including RNA splicing, protein translation, proteostasis imbalance and nuclear export as TDP-43 phase regulators. Importantly, TDP-43 phase transition can be dynamically recapitulated in vitro in a semi-permeabilized cell system, which reveals that the inhibition of nuclear export reshapes the nuclear environment in favor of an RNA-dependent TDP-43 liquid-liquid phase separation (LLPS) state, which mitigates cytoplasmic TDP-43 aggregation. We validated this mechanism in a brain organoid model bearing an ALS-associated mutation, showing that nuclear export deficiency can limit pathogenic phospho-TDP-43 accumulation. These findings establish nuclear export as a key regulator of TDP-43 phase transitions and define a mechanistic framework that links altered nuclear transport and phase dynamics to TDP-43 aggregation potential.

Revised version: RNA-binding protein TAR DNA-binding protein 43 (TDP-43) can form liquid-like, nuclear assemblies whose phase behavior is thought to influence its aggregation propensity and neurotoxic activity. The mechanism(s) that governs the liquid-to-solid phase transition of TDP-43 is poorly defined. Here we combined chemical and genome-wide genetic screenings to identify cellular factors that modulate the phase behavior of an RNA-binding defective TDP-43 mutant. Our screens uncovered multiple cellular processes including RNA splicing, protein translation, proteostasis imbalance and nuclear export as TDP-43 phase regulators. We also developed a semi-permeabilized cell system that partially recapitulates the TDP-43 phase transition in vitro, and showed that nuclear export inhibition reshapes the nuclear environment to favor RNA-dependent liquid-liquid phase separation (LLPS) of TDP-43, which mitigates its aggregation. Nuclear export inhibition in a brain organoid model bearing an ALS-associated mutation reduces pathogenic phospho-TDP-43 accumulation. These findings identify multiple modulators of TDP-43 phase transitions in a sensitized model system and establish a framework for further dissecting the link between nuclear transport and TDP-43 phase dynamics.

Added in the revised version: Introduction

Removed in the revised version: The RNA-binding protein TAR DNA-binding protein 43 (TDP-43), encoded by the TARDBP gene, plays a central role in RNA transport, splicing and metabolism 1, 2, and is the major pathological component of cytoplasmic inclusions in amyotrophic lateral sclerosis (ALS) and frontotemporal dementia (FTD) 3–5.

Earlier version: TDP-43 contains two RNA-recognition motifs (RRMs), a bipartite nuclear-localization sequence, and a glycine-rich C-terminal low-complexity domain (LCD) that is intrinsically prone to aggregation 6. Under physiological conditions, TDP-43 binds RNAs in the nucleus, participating in multiple RNA-metabolic processes. In contrast, disease-causing mutations trigger its mis-localization to the cytoplasm, forming detergent-resistant inclusions. TDP-43-positive protein aggregates were present in ∼97% of ALS and ∼45% of FTD cases, as well as in a subset of Alzheimer’s disease patients. These diseases are now collectively termed TDP-43 proteinopathies 4, 7. Both loss of nuclear function and gain of cytoplasmic toxicity have been implicated in neurodegeneration 2, 6, underscoring the importance of nucleocytoplasmic transport in regulating TDP-43 patho-physiological functions.

Revised version: The RNA-binding protein TAR DNA-binding protein 43 (TDP-43), encoded by the TARDBP gene, plays a central role in RNA transport, splicing and metabolism (Guo and Shorter, 2017; Suk and Rousseaux, 2020), and is the major pathological component of cytoplasmic inclusions in amyotrophic lateral sclerosis (ALS) and frontotemporal dementia (FTD) (Kellett et al., 2025; Neumann et al., 2006; Rummens and Da Cruz, 2025). TDP-43 contains two RNA-recognition motifs (RRMs), a bipartite nuclear-localization sequence, and a glycine-rich C-terminal low-complexity domain (LCD) that is intrinsically prone to aggregation (Tziortzouda et al., 2021). Under physiological conditions, TDP-43 binds RNAs in the nucleus, participating in multiple RNA-metabolic processes. In contrast, disease-causing mutations trigger its mis-localization to the cytoplasm, forming detergent-resistant inclusions. TDP-43-positive protein aggregates were present in ∼97% of ALS and ∼45% of FTD cases, as well as in a subset of Alzheimer’s disease patients. These diseases are now collectively termed TDP-43 proteinopathies (Brettschneider et al., 2014; Neumann et al., 2006). Both loss of nuclear function and gain of cytoplasmic toxicity have been implicated in neurodegeneration (Suk and Rousseaux, 2020; Tziortzouda et al., 2021), underscoring the importance of nucleocytoplasmic transport in regulating TDP-43 patho-physiological functions.

Earlier version: Like many RNA-binding proteins (RBPs), TDP-43 undergoes liquid-liquid phase separation to form membraneless biomolecular condensates 6, 8, 9. Under conditions of cellular stress or in the presence of RNA-binding inhibiting mutations, these condensates can transition from dynamic liquid droplets to rigid gel-like or solid assemblies 8, 10, 11. To date, more than 70 disease-linked mutations in TARDBP have been identified in ALS or FTD, many within the LCD, while others clustering near or in the RRMs to disrupt RNA binding 6. Defects in RNA binding appears to be a major determinant of TDP-43 phase behavior 12, 13. Disease-associated mutations diminish TDP-43’s activities in RNA metabolism, causing widespread cryptic exon inclusion 14–17, alternative polyadenylation and other RNA maturation defects that collectively contribute to neurotoxicity 18, 19.

Revised version: Like many RNA-binding proteins (RBPs), TDP-43 undergoes liquid-liquid phase separation to form membraneless biomolecular condensates (Conicella et al., 2016; Shorter, 2019; Tziortzouda et al., 2021). Under conditions of cellular stress or in the presence of RNA-binding mutations, these condensates can transition from dynamic liquid droplets to rigid gel-like or solid assemblies (Conicella et al., 2016; Gasset-Rosa et al., 2019; Lang et al., 2024). To date, more than 70 disease-linked mutations in TARDBP have been identified in ALS or FTD, many within LCD, while others clustering near or in RRMs to disrupt RNA binding (Tziortzouda et al., 2021). Defects in RNA binding appears to be a major determinant of TDP-43 phase behavior (Babinchak et al., 2019; Mann et al., 2019). Disease-associated mutations diminish TDP-43’s activities in RNA metabolism, causing widespread cryptic exon inclusion (Ling et al., 2015; Ma et al., 2022; Mehta et al., 2023; Zhang et al., 2025b), alternative polyadenylation and other RNA maturation defects that collectively contribute to neurotoxicity (Bryce-Smith et al., 2025; Zeng et al., 2025).

Earlier version: Apart from disease-associated mutations, post-translational modifications (PTM), especially acetylation, have been shown to modulate TDP-43’s RNA binding capacity 16, 20. Acetylation was detected on endogenous TDP-43 in ALS patient samples. In vitro, TDP-43 acetylation is enhanced when cells are exposed to oxidative or proteotoxic stress 20, 21. Interestingly, like disease-associated TDP-43 mutants, acetylated TDP-43 is also more aggregation-prone and displays a similar phase regulation pattern 20. For example, a recent study showed that a TDP-43 acetylation mimetic mutant designated as 2KQ and several RNA-binding defective TDP-43 disease variants can all undertake a unique form of demixing, forming “anisosomes” that contain an anisotropic spherical shell of TDP-43 and a central liquid core enriched in the HSP70 chaperone 22. Anisosomes, observed mostly in the nucleus under overexpression conditions, have not been validated in human patient samples. Nevertheless, they might be an intermediate en route to pathological aggregation when cellular proteostasis or nucleocytoplasmic transport is compromised, as suggested by a recent study 22. Additional in vitro studies using purified TDP-43 and small molecules have begun to elucidate the molecular basis of TDP-43 self-association 23, 24, offering insights into its phase behavior and aggregation mechanisms.

Revised version: Apart from disease-associated mutations, post-translational modifications (PTM), especially acetylation, have been shown to modulate TDP-43’s RNA binding capacity (Cohen et al., 2015; Zhang et al., 2025b). Acetylation was detected on endogenous TDP-43 in ALS patient samples. Although the abundance of acetylated TDP-43 has not been determined, acetylated TDP-43 is enhanced in cells exposed to oxidative or proteotoxic stress (Cohen et al., 2015; Wang et al., 2017). Interestingly, like disease-associated TDP-43 mutants, acetylated TDP-43 is also more aggregation-prone and displays a similar phase regulation pattern (Cohen et al., 2015). Moreover, a recent study showed that a TDP-43 acetylation mimetic mutant designated as 2KQ behaves similarly as RNA-binding defective TDP-43 variants as they all undergo a unique form of demixing, forming “anisosomes” that contain an anisotropic spherical shell of TDP-43 and a central liquid core enriched in the HSP70 chaperone (Yu et al., 2021).

Added in the revised version: Anisosomes were observed mostly in the nucleus under overexpression conditions. Although their existence in humans and relevance to ALS or FTD are unclear, TDP43 2KQ offers a platform for dissecting the potential link between TDP-43 phase behaviors (droplet formation, gel structures etc.) and aggregation (Yu et al., 2021).

Earlier version: Despite these advances, the cellular pathways modulating TDP-43 phase behavior (droplet formation, anisosomal shell–core architecture, conversion to gel/solid structures) and its link to disease-associated aggregation remain unclear. We address this knowledge gap by combining a chemical-genetic screen approach with a genome-wide siRNA screen to identify molecular determinants that modulate the phase behavior of the RNA-binding defective TDP-43 2KQ mutant. Our study reveals critical contributions of RNA splicing, protein translation, and the HSP90/ubiquitin–proteasome proteostasis network to TDP-43 phase transition. Importantly, our work demonstrates how nuclear export can influence the transition of TDP-43 from a nuclear liquid-demixed form to a cytoplasmic immobile gel-like structure akin to protein aggregation. Using an iPSC-derived 3-D organoid model bearing an ALS-associated mutation 16, we further validated the impact of nuclear export inhibition on the accumulation of cytoplasmic phosphorylated-TDP-43. Collectively, these findings establish a mechanistic framework linking altered phase dynamics and TDP-43 aggregation to nuclear transport defects, a process known to modulate neurodegeneration in ALS and FTD.

Revised version: In this study, we combined a chemical-genetic screen approach with a genome-wide siRNA screen to identify molecules that modulate the phase behavior of TDP-43 2KQ. Our study reveals critical contributions of RNA splicing, protein translation, and the HSP90/ubiquitin–proteasome proteostasis network to TDP-43 phase transition. Additionally, nuclear export inhibition can influence the transition of TDP-43 from a nuclear liquid-demixed form to a cytoplasm-localized immobile gel-like structure akin to protein aggregation. In an iPSC-derived organoid model bearing an ALS-associated mutation (Zhang et al., 2025b), nuclear export inhibition reduces cytoplasmic phosphorylated-TDP-43. Collectively, these findings link TDP-43 phase dynamics and aggregation to nuclear transport and other cellular processes.

Results

Earlier version: A chemical genetic screen identified modulators of TDP-43 phase behavior

Revised version: A chemical genetic screen identified modulators of TDP-43 phase separation

Earlier version: To identify cellular pathways modulating TDP-43 phase behavior, we conducted a chemical genetic screen using a small molecule library targeting diverse known cellular pathways and processes. We employed a previously established DLD1 cell model stably expressing an RNA-binding defective TDP-43 mutant (2KQ) tagged with the green fluorescence protein Clover. In this system, TDP-43 2KQ undergoes rapid demixing upon induced expression 22, generating 20-30 sphere objects named anisosome in each nucleus. Both the number and size of anisosome increase substantially in the initial phase of TDP-43 induction but remain largely unchanged after 24 h of induction. We seeded cells in 384-well plates and induced the expression of TDP-43 2KQ with doxycycline for 24 hours. We then treated these cells with a LOPAC library containing 1280 drugs for 24 hours (Figure 1A). This approach could avoid false positive drugs that affect the expression of TDP-43 2KQ.

Revised version: To identify cellular pathways modulating TDP-43 phase behavior, we conducted a chemical genetic screen using a small molecule library targeting diverse known cellular pathways and processes. We employed a previously established DLD1 cell model stably expressing an RNA-binding defective TDP-43 mutant (2KQ) tagged with the green fluorescence protein Clover, because this mutant undergoes rapid demixing upon induced expression (Yu et al., 2021), generating 20-30 sphere objects named anisosome in each nucleus. Both the number and size of anisosome increase substantially in the initial phase of TDP-43 induction but remain largely unchanged after 24 h of induction. We seeded cells in 384-well plates and induced the expression of TDP-43 2KQ with doxycycline for 24 hours. We then treated these cells with a LOPAC library containing 1280 drugs for 24 hours (Figure 1A). This approach could avoid false positive drugs that affect the expression of TDP-43 2KQ.

A chemical genetic screen identifies TDP-43 phase modulators.

Earlier version: (A) Workflow of the chemical genetic screen. (B) Dose dependent reduction of anisosome number by identified chemicals. DLD1 TDP-43 2KQ-Clover cells were treated with doxycycline to induce anisosome for 24 h and then treated with the drugs as indicated for 15 h. Cells were imaged for anisosome count.

Revised version: (A) Workflow of the chemical genetic screen. (B) Dose dependent reduction of anisosome number by identified chemicals. DLD1 TDP-43 2KQ-Clover cells were treated with doxycycline to induce anisosome for 24 h and then treated with drugs at the indicated concentrations. At least 500 cells were imaged for anisosome count per condition. A non-linear regression model was used to fit the data by GraphPad Prism.

Earlier version: High content imaging combined with algorithm-based image analyses (see methods) detected ∼20 spherical anisosomes per cell in controls and most treated cells. However, in cells treated with a subset of drugs, the number of anisosomes was noticeably reduced (Supplemental Table 1). Importantly, cells treated with Bortezomib, a potent proteasome inhibitor, had fewer but enlarged TDP-43 positive puncta with irregular shapes (Figure 1B) similarly to cells treated with another proteasome inhibitor as reported previously 22, 25. This result validated our screening design.

Revised version: High content imaging combined with algorithm-based image analyses (see methods) detected ∼20 spherical anisosomes per cell in controls and most treated cells. However, in cells treated with a subset of the drugs, the number of anisosomes was noticeably reduced (Supplemental Table 1). Importantly, cells treated with Bortezomib, a potent proteasome inhibitor, had fewer but enlarged TDP-43 positive puncta with irregular shapes (Figure 1B) similarly to cells treated with another proteasome inhibitor (van Eersel et al., 2011; Yu et al., 2021). This result validated our screening design.

Earlier version: To further confirm the identified hits, we performed a concentration titration experiment, which showed that these inhibitors dose-dependently reduced the anisosome number with IC50 ranging from sub-micromolar to micromolar (Figure 1B, Supplemental Table 1). Immunoblotting showed that treatment with these drugs did not change the total TCP-43 protein level (Supplementary Figure 1), suggesting that the impact on anisosome is not due to altered protein expression. Interestingly, the screen identified cytoplasmic signaling kinases including PAK4, BTK, and GSK-3β as well as the receptor tyrosine kinase PDGFRβ as modulators of TDP-43 phase behavior (Figure 1B, Supplemental Table 1). PAK4 and GSK-3β are known regulators of the YAP signaling pathway 26, and YAP was recently identified as a TDP-43-interacting protein that influences its phase dynamics 27. These connections provide additional physiological support for the relevance of the uncovered modulators. Our screen also linked TDP-43 2KQ demixing to protein translation because cycloheximide, a ribosome elongation inhibitor also affect TDP-43 phase dynamics (see below).

Revised version: To further confirm the identified hits, we performed a concentration titration experiment, which showed that these inhibitors dose-dependently reduced the anisosome number with IC50 ranging from sub-micromolar to micromolar (Figure 1B, Supplemental Table 1). Immunoblotting showed that treatment with these drugs did not change the total TDP-43 protein level (Supplementary Figure 1), suggesting that the effect on anisosome is not due to altered protein expression. Interestingly, the screen identified cytoplasmic signaling kinases including PAK4, BTK, and GSK-3β as well as the receptor tyrosine kinase PDGFRβ as modulators of TDP-43 phase behavior (Figure 1B, Supplemental Table 1). PAK4 and GSK-3β are known regulators of the YAP signaling pathway (Jin et al., 2025), and YAP was recently identified as a TDP-43-interacting protein that influences its phase dynamics (Zhang et al., 2025a). These connections provide additional physiological validation for the uncovered modulators. Our screen also links TDP-43 2KQ demixing to protein translation regulation because cycloheximide, a ribosome elongation inhibitor also affects TDP-43 phase dynamics (see below).

Live cell imaging revealed two distinct classes of TDP-43 phase modulators

Earlier version: We next used confocal microscopy to further characterize a subset of the identified inhibitors. We chose Tripterin (a drug targeting the heat shock protein HSP90), CP-673451 (a PDGFRβ inhibitor), two deubiquitinase inhibitors VLX-1570 and Spautin-1, and the Exportin-1 inhibitor KPT-276 because previous reports have linked the targets of these compounds to TDP-43 proteinopathies 28–31. We first induced anisosome formation for 24 hours and then treated cells with the inhibitors. Compared to DMSO-treated controls, cells treated with Spautin-1 had fewer and smaller anisosomes. KPT-276 treatment also reduced anisosome number, albeit to a lesser extent (Figure 2A, top panels). In contrast, cells treated with CP-673451, Tripterin, or Bortezomib had fewer but larger TDP-43 punctae. Since Tripterin was also reported as a 20S proteasome inhibitor, we confirmed the role of HSP90 in this process using the well-established HSP90 inhibitor Geldanamycin (Supplemental Figure 2A). These results suggest that TDP-43 phase dynamics are modulated by the cellular folding capacity, the ubiquitin proteasome system, and nucleocytoplasmic transport.

Revised version: We next used confocal microscopy to further characterize selected inhibitors. We chose Tripterin (a drug targeting the heat shock protein HSP90), CP-673451 (a PDGFRβ inhibitor), two deubiquitinase inhibitors VLX-1570 and Spautin-1, and the Exportin-1 inhibitor KPT-276 because previous reports have linked the targets of these compounds to TDP-43 proteinopathies (Archbold et al., 2018; Farrawell et al., 2020; Furukawa et al., 2015; Garcia-Toscano et al., 2024). We first induced anisosome formation for 24 hours and then treated cells with these inhibitors. Compared to DMSO-treated controls, cells treated with Spautin-1 had fewer and smaller anisosomes. KPT-276 treatment also reduced anisosome number, albeit to a lesser extent (Figure 2A, top panels). In contrast, cells treated with CP-673451, Tripterin, or Bortezomib had fewer but larger TDP-43 puncta. Since Tripterin was also reported as a 20S proteasome inhibitor, we confirmed the involvement of HSP90 in this process using the well-established HSP90 inhibitor Geldanamycin (Supplemental Figure 2A). These results suggest that TDP-43 phase dynamics is modulated by the cellular folding capacity, the ubiquitin proteasome system, and nucleocytoplasmic transport.

Two distinct types of TDP-43 phase modulators.

Earlier version: (A) DLD1 TDP-43 2KQ-Clover cells treated with doxycycline for 24 h were further treated with the indicated compounds for 6 h and imaged (SPA, Spautin-1, 5 μM; KPT, KPT-276, 15 μM; CP, CP-673451, 30 μM; TRP, Tripterin, 1 μM; BTZ, Bortezomib, 10 nM). Shown are single confocal z-section views. Scale bar, 5 μm. (B) FRAP analyses of anisosomes after treatment with the indicated drugs for 3-5 h. Circles indicate bleached areas. Shown are single confocal z-section views. Scale bars, 1 μm. (C, D) Quantification of the experiments represented in B. MFI, Mean Fluorescence Intensity. N = 5 anisosomes/condition. (E) Reverse FRAP analyses of anisosome dynamics. The areas indicated by the dashed lines were bleached. Shown are single confocal z-section views. Scale bar, 5 μm. (F) Quantification of the Mean Fluorescence Intensity (MFI) loss over time in unbleached anisosomes as shown in E. N = 7 anisosomes/condition. (G) Quantification of the initial rate of fluorescence loss in unbleached anisosomes as shown in E. ****, p<0.0001; ***, p<0.001 by unpaired Student’s t-test (two-tailed). N = 7 anisosomes/condition.

Revised version: (A) DLD1 TDP-43 2KQ-Clover cells treated with doxycycline for 24 h were further treated with the indicated compounds for 6 h and imaged (SPA, Spautin-1, 5 μM; KPT, KPT-276, 15 μM; CP, CP-673451, 30 μM; TRP, Tripterin, 1 μM; BTZ, Bortezomib, 10 nM). Shown are single confocal z-section views. Scale bar, 5 μm. (B) FRAP analyses of anisosomes after treatment with the indicated drugs for 3-5 h. Circles indicate bleached areas. Shown are single confocal z-section views. Scale bars, 1 μm. (C, D) Quantification of the experiments represented in B. MFI, mean fluorescence intensity. N = 5 anisosomes/condition. (E) Reverse FRAP analyses of anisosome dynamics. The areas indicated by the dashed lines were bleached. Shown are single confocal z-section views. Scale bar, 5 μm. (F) Quantification of the mean fluorescence intensity (MFI) loss over time in unbleached anisosomes as shown in E. 4-6 anisosomes were analyzed for each condition. (G) Quantification of the initial rate of fluorescence loss of unbleached anisosomes as shown in E. ****, p<0.0001; ***, p<0.001 by unpaired Student’s t-test (two-tailed). The rate of fluorescence loss in the first 10 seconds after photobleaching was calculated. 4-6 anisosomes were analyzed for each condition.

Earlier version: To further dissect the effects of these inhibitors on TDP-43 phase separation, we combined drug treatment with live cell fluorescence imaging, using Fluorescence Recovery After Photobleaching (FRAP) or reversed FRAP. In FRAP, when the green fluorescence of anisosomes was bleached by a laser in untreated cells, it rapidly recovered due to the recruitment of TDP-43 2KQ from the nucleoplasm (Figure 2B, top panels, Figure 2C, D). Conversely, in reversed FRAP, bleaching TDP-43 2KQ surrounding an anisosome repeatedly over time caused a gradual decline in fluorescence intensity within the unbleached anisosome, while the surrounding areas regained fluorescence partially (Figure 2E, top panel). This result suggests that TDP-43 2KQ in anisosomes undergoes rapid exchange with a pool in the nucleoplasm. Likewise, in cells treated with Spautin-1 or KPT-276, anisosome-associated TDP-43 2KQ could freely exchange with the nucleoplasmic pool (Figure 2B, D). In contrast, in cells treated with Bortezomib, Tripterin or Geldanamycin, the enlarged TDP-43 puncta appeared static; TDP-43 2KQ fluorescence within anisosomes failed to recover after photobleaching (Figure 2B, C, Supplemental Figure 2B), nor did it decrease over time when surrounding TDP-43 2KQ was bleached (Figure 2E-G). Collectively, these results suggest that these inhibitors modulate anisosomes via two distinct mechanisms: one maintains TDP-43 in a demixed liquid state, while the other converts it to a gel-like solid state.

Revised version: To further dissect the effects of these inhibitors on TDP-43 phase separation, we combined drug treatment with live cell fluorescence imaging, using Fluorescence Recovery After Photobleaching (FRAP) or reversed FRAP. In FRAP, when the green fluorescence of anisosomes was bleached by a laser in untreated cells, it rapidly recovered due to the recruitment of TDP-43 2KQ from the nucleoplasm (Figure 2B, top panels, Figure 2C, D). Conversely, in reversed FRAP, bleaching TDP-43 2KQ surrounding an anisosome repeatedly over time caused a gradual decline in fluorescence intensity within the unbleached anisosome, while the surrounding areas regained fluorescence partially (Figure 2E, top panels). This result suggests that TDP-43 2KQ in anisosomes undergoes rapid exchange with a pool in the nucleoplasm. Likewise, in cells treated with Spautin-1 or KPT-276, anisosome-associated TDP-43 2KQ could freely exchange with the nucleoplasmic pool (Figure 2B, D). In contrast, in cells treated with Bortezomib, Tripterin or Geldanamycin, the enlarged TDP-43 puncta appeared static; TDP-43 2KQ fluorescence within anisosomes failed to recover after photobleaching (Figure 2B, C, Supplemental Figure 2B), nor did it decrease over time when surrounding TDP-43 2KQ was bleached (Figure 2E-G). Collectively, these results suggest that these inhibitors modulate anisosomes via two distinct mechanisms: some maintain TDP-43 in a demixed liquid state, while others convert anisosomes to a gel-like solid state.

A genome-wide siRNA screen identified genetic modulators of TDP-43 phase separation

Earlier version: To further decipher the molecular determinants of TDP-43 phase separation, we conducted an unbiased genome-wide siRNA knockdown (KD) screen (Figure 3A). To this end, we transfected a siRNA library targeting 21,404 human genes each with three siRNAs into DLD1 cells stably expressing Clover-tagged TDP-43 2KQ. After gene KD, anisosomes were induced for 24 hours. High content imaging and automated analysis identified 1,533 candidates whose KD reduced the anisosome number per cell (Z-score > 2) (Supplemental Table S2). To further narrow down the list, we performed a STRING protein network analysis based on the assumption that a protein interaction network bearing multiple positive hits would be more likely to be a true effector. We selected 211 networked genes with top Z-scores and re-screened each of them with three additional siRNAs. This suggested the involvement of 110 genes in TDP-43 anisosome regulation (Figure 3A, Supplementary Table S3). GO pathway analyses categorized these genes into several pathways including RNA splicing, protein translation, proteasomal degradation, and nuclear transport (Figure 3B). The identification of ribosomal proteins and proteasome subunits is consistent with our chemical genetic screen, which implicates translation and proteasomal degradation in anisosome modulation (Supplemental Table 1). Interestingly, GO analyses using ‘molecular function’ linked many identified genes to neurodegenerative diseases, particularly ALS (Figure 3C, D).

Revised version: To further decipher the molecular determinants of TDP-43 phase dynamics, we conducted an unbiased genome-wide siRNA knockdown (KD) screen (Figure 3A). To this end, we transfected a siRNA library targeting 21,404 human genes each with three siRNAs into DLD1 cells stably expressing Clover-tagged TDP-43 2KQ. After gene KD, anisosomes were induced for 24 hours. High content imaging and automated analysis identified 1,533 candidates whose KD reduced the anisosome number per cell (Z-score > 2) (Supplemental Table S2). To further narrow down the list, we performed a STRING protein network analysis based on the assumption that a protein interaction network bearing multiple positive hits would be more likely a true effector. We selected 211 networked genes with top Z-scores and re-screened each of them with three additional siRNAs. The result suggested the involvement of 110 genes in TDP-43 anisosome regulation (Figure 3A, Supplementary Table S3). GO pathway analyses categorized these genes into several pathways including RNA splicing, protein translation, proteasomal degradation, and nuclear transport (Figure 3B). The identification of ribosomal proteins and proteasome subunits is consistent with our chemical genetic screen, which has implicated translation and proteasomal degradation in anisosome modulation (Supplemental Table 1). Interestingly, GO analyses using ‘molecular function’ linked many identified genes to neurodegenerative diseases, particularly ALS (Figure 3C, D).

A genome-wide siRNA screen identifies modifiers of TDP43 phase behavior.

Earlier version: (A) The workflow of the siRNA genetic screen. GW, genome-wide; AS, anisosome; HT, high-throughput. (B) Pathway analysis of genes whose knockdown reduces anisosome number in cells. The relative anisosome count are indicated by both color and size of the nodes. (C) GO molecular function analysis of TDP-43 phase modifiers. (D) A list of identified genes linked to ALS in C.

Revised version: (A) The workflow of the siRNA genetic screen. GW, genome-wide; AS, anisosome; HT, high-throughput. (B) Pathway analysis of genes whose knockdown reduces anisosome number in cells. The relative anisosome count are indicated by both color and size of the nodes. Candidate genes were selected based on Z-score >2 after normalization to the siRNA negative control. (C) GO molecular function analysis of TDP-43 phase modifiers. (D) A list of identified genes linked to ALS in C.

Anisosome dynamics is modulated by RNA splicing and protein translation

Given the well-established function of TDP-43 in RNA binding and processing, we investigated the role of RNA splicing in anisosome regulation. To this end, we first incubated TDP-43 2KQ cells with doxycycline to induce anisosomes and then treated cells with a potent RNA splicing inhibitor Pladienolide-B (PlaB). Confocal microscopy confirmed that RNA splicing inhibition resulted in fewer but larger TDP-43-positive puncta in a dose dependent manner (Figure 4A, B, Supplemental Figure 3A).

Anisosome phase behavior is modulated by RNA splicing and protein translation.

Earlier version: (A) The splicing inhibitor Pladienolide-B (PlaB) reduces anisosome number in a dose dependent manner. Representative images of anisosome-induced cells treated with DMSO (control), 5 nM, or 20 nM PlaB for 16 h. Shown are maximum intensity projection views. Scale bar: 10 µm. (B) Quantification showing the number of anisosome (AS) per cell in TDP-43 2KQ-Clover cells treated with PlaB as indicated. * p <0.05, ** p <0.01, **** p < 0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. N=3 biological repeats, and in each experiment, at least 30 randomly selected cells were analyzed. (C) Representative z-section confocal FRAP images of anisosomes in cells treated for 16 h with DMSO or PlaB (20 nM). BB, before photobleaching, right after photobleaching (0 s), or 4 and 30 seconds after photobleaching (4 s and 30 s). Scale bar, 1 µm. (D) The graph shows the quantification of the remaining TDP-43 fluorescence (FL) in C. Error bars indicate mean ±SD, N = 28 for control and 23 for PlaB-treated cells. MFI, Mean Fluorescence Instensity. (E) Live cell imaging of anisosome fusion in TDP-43 2KQ-Clover cells treated with 20 nM PlaB for 5 h before tracking the fusion. Representative reconstructed 3-D images from Movie S1 showing fusion events indicated by arrows. Scale bar, 1 µm. (F) Representative maximum intensity projection views of anisosome-induced (24 h) cells treated with DMSO (control) or ANS (200 nM) for 16 h. Scale bar, 1 µm. (G) Quantification of the number of anisosomes per cell in randomly selected images of DMSO-, Cycloheximide (CHX)-, or ANS-treated cells. **** p < 0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. Each dot represents a randomly selected field with at least 20 cells counted from one of the 3 biological repeats. (H) Quantification of anisosome size in control or cells treated with CHX or ANS. ****, p < 0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. N=3 biological repeats. AU, arbitrary unit. (I) As in C except that anisosome-induced cells were treated with CHX or ANS before photobleaching. Scale bar 1 µm. (J) Quantification of fluorescence recovery in CHX- or ANS-treated cells vs the DMSO control. N=10 anisosomes/condition.

Revised version: (A) The splicing inhibitor Pladienolide-B (PlaB) reduces anisosome number in a dose dependent manner. Representative images of anisosome-induced cells treated with DMSO (control), 5 nM, or 20 nM PlaB for 16 h. Shown are maximum intensity projection views. Scale bar: 10 µm. (B) Quantification showing the number of anisosome (AS) per cell in TDP-43 2KQ-Clover cells treated with PlaB as indicated. * p <0.05, ** p <0.01, **** p < 0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. n=3 independent experiments, and in each experiment, at least 30 randomly selected cells were analyzed. (C) Representative z-section confocal FRAP images of anisosomes in cells treated for 16 h with DMSO or PlaB (20 nM). BB, before photobleaching, right after photobleaching (0 s), or 4 and 30 seconds after photobleaching (4 s and 30 s). Scale bar, 1 µm. (D) The graph shows the quantification of the remaining TDP-43 fluorescence (FL) in C. Error bars indicate mean ± SD, N = 28 for control and 23 for PlaB-treated cells. MFI, mean fluorescence instensity. (E) Live cell imaging of anisosome fusion in TDP-43 2KQ-Clover cells treated with 20 nM PlaB for 5 h before tracking the fusion. Representative reconstructed 3-D images from Movie S1 showing fusion events indicated by arrows. Scale bar, 1 µm. (F) Representative maximum intensity projection views of anisosome-induced (24 h) cells treated with DMSO (control) or ANS (200 nM) for 16 h. Scale bar, 1 µm. (G) Quantification of the number of anisosomes per cell in randomly selected images of DMSO-, Cycloheximide (CHX)-, or ANS-treated cells. **** p < 0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. Each dot represents a randomly selected field with at least 20 cells counted from one of the 3 independent experiments. (H) Quantification of anisosome size in control or cells treated with CHX or ANS. ****, p < 0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. N=3 independent experiments. AU, arbitrary unit. (I) As in C except that anisosome-induced cells were treated with CHX or ANS before photobleaching. Scale bar 1 µm. (J) Quantification of fluorescence recovery in CHX- or ANS-treated cells vs the DMSO control. N=10 anisosomes/condition.

Despite size difference, anisosomes in drug-treated cells were morphologically indistinguishable from controls. FRAP experiments further showed that TDP-43 within PlaB-treated condensates remained highly mobile, although the fluorescence recovery rate was slightly reduced compared to controls (Figure 4C, D). Time-lapse live cell imaging frequently detected fusion of TDP-43 puncta after PlaB treatment (Figure 4E, Movie S1). Together, these results suggest that disrupting RNA splicing stabilizes TDP-43 in a demixed liquid state, resulting in larger condensates.

Next, we explored the effect of translation inhibition on anisosome dynamics using Cycloheximide (CHX) and Anisomycin (ANS), two translation elongation inhibitors. Confocal imaging and immunoblotting showed that translation inhibition by CHX or ANS did not significantly reduce TDP-43 expression (Supplemental Figure 1, 3). Like PlaB-treated cells, CHX- and ANS-treated cells contained fewer, enlarged anisosomes (Figure 4F-H, Supplemental Figure S2B). FRAP analysis revealed that TDP-43 dynamics within anisosomes were similar between ANS-treated and control cells (Figure 4I, J). These results suggest that translation inhibition also stabilizes TDP-43 in liquid condensates.

Nuclear export modulates TDP-43 liquid-to-solid phase transition

Earlier version: Our screen also identified several nuclear transport regulators (e.g., Exportin-1/XPO1, Exportin-2/CSE1L) and nuclear pore components as anisosome regulators (Figure 3B). We focused on XPO1 because our chemical genetic screen identified two XPO1 inhibitors, KPT-276 and Verdinexor as potent anisosome modulators (Figure 1B).

Revised version: Our screen also identified several nuclear transport regulators (e.g., Exportin-1/XPO1, Exportin-2/CSE1L) and nuclear pore components as anisosome modulators (Figure 3B). We focused on XPO1 because our chemical genetic screen identified two XPO1 inhibitors, KPT-276 and Verdinexor as potent anisosome modulators (Figure 1B).

Earlier version: To explore the role of XPO1 in TDP-43 anisosome regulation, we induced anisosome formation in Clover-TDP-43 2KQ cells and treated them with Leptomycin B (LMB), a potent XPO1 inhibitor. LMB treatment resulted in a time-and dose-dependent reduction in anisosome number, while increasing their sizes (Figure 5A-C). The enlarged TDP-43 puncta showed the typical hollowed anisosome ring structure (Figure 5D), indicating that TDP-43 remained in liquid state. Indeed, FRAP experiments demonstrated that LMB treatment did not significantly affect the fluorescence recovery rate of TDP-43 within bleached anisosomes (Figure 5D, E). Time-lapse confocal imaging showed that shortly after LMB treatment, preformed anisosomes began to fuse with each other (Figure 5F, Movie S2). Thus, XPO1 inhibition stabilizes TDP-43 in a liquid phase.

Revised version: To explore the role of XPO1 in TDP-43 anisosome regulation, we induced anisosome formation in Clover-TDP-43 2KQ cells and treated them with Leptomycin B (LMB), a potent XPO1 inhibitor. LMB treatment resulted in a time- and dose-dependent reduction in anisosome number, while increasing their sizes (Figure 5A-C). The enlarged TDP-43 puncta showed the typical hollowed anisosome ring structure (Figure 5D), indicating that TDP-43 remained in a demixed liquid state. Indeed, FRAP experiments demonstrated that LMB treatment did not significantly affect the fluorescence recovery rate of TDP-43 within bleached anisosomes (Figure 5D, E). Time-lapse confocal imaging showed that preformed anisosomes began to fuse with each other 5 hours after LMB treatment (Figure 5F, Movie S2). Thus, XPO1 inhibition stabilizes TDP-43 in a liquid phase.

XPO1 regulates anisosome liquid-to-solid transition.

Earlier version: (A) Pharmacological inhibition of XPO-1 with Leptomycin B (LMB, 200 nM) reduces the number of anisosome. The graph on top indicates the experimental design. Arrows show cells with enlarged anisosomes. Scale bar 10 µm (B) Quantification of anisosome numbers per cell in experiments represented by A. Error bars indicate mean ±SD; **** p <0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. N=3 biological repeats. (C) LMB dose dependently reduces anisosome number. Anisosome were induced in TDP-43 2KQ-Clover cells followed by treatment with LMB at the indicated concentrations for 16 h. The histogram shows the distribution of cells (50-90 cells/condition) by anisosome count. (D) FRAP experiments demonstrate that anisosomes remain in a liquid phase following LMB treatment (200 nM, 16 h). Anisosome-induced (24 h) cells were treated with DMSO or LMB for 16 h and then photobleached at the indicated areas. Scale bar 1 µm. (E) Quantification of the FRAP experiment in D. N=18 for control and 16 anisosomes for LMB-treated cells. (F) Time-lapse confocal microscopy detects anisosome fusion after LMB (200 nM, 5 h) treatment in TDP-43 2KQ-Clover cells. Arrows indicate fusion events. Scale bar, 1 µm (G) Overexpression of mCherry-tagged XPO-1 in TDP-43 2KQ-Clover cells induces cytoplasmic TDP-43 puncta. TDP-43 2KQ-Clover cells (green) transfected with mCherry-XPO1 (red) were stained with DAPI (blue) to label nuclei (dashed lines). Cells were imaged 48 h post-transfection. Panels 1, 2 show a representative confocal section, while panels 3-6 show reconstructed 3-D views. The position and volume of anisosomes were also presented in magenta in a surface-rendered view (SRV) in panel 5. The arrow in panel 3 indicates an example of cytoplasmic TDP-43 aggregate. Scale bar, 10 µm. (H) Quantification of the percentage of cells showing cytosolic TDP-43 puncta in XPO-1 positive (pos) and negative (neg) cells in randomly selected fields from 3 independent experiments. **** p <0.0001 by two-tailed unpaired Student’s t-test. (I) FRAP-based confocal imaging reveals the transition of anisosomes into a gel-like state upon XPO-1 overexpression. Scale bar, 1 µm. (J) Anisosome formation changes endogenous XPO-1 localization. TDP-43 2KQ-Clover before or after anisosome induction were stained with anti-XPO-1 antibodies (red) and DAPI (blue). The bleached areas were indicated by dashed lines. Scale bar, 10 µm. (K) Quantification of nuclear endogenous XPO-1 Mean Fluorescence Intensity (MPI) in individual cells (indicated by dots) before or after anisosome induction. **** p < 0.001 by two-tailed unpaired Student’s t-test. N = 58 cells for uninduced and 63 cells for induced condition in 2 independent repeats. (L) A schematic model depicting how the nuclear XPO-1 activity influences TDP-43 liquid-to-phase transition. AS, anisosome.

Revised version: (A) Pharmacological inhibition of XPO-1 with Leptomycin B (LMB, 200 nM) reduces the number of anisosome. The graph on top indicates the experimental design. Arrows show cells with enlarged anisosomes. Scale bar 10 µm (B) Quantification of anisosome numbers per cell in experiments represented by A. Error bars indicate mean ± SD; **** p <0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. n=3 independent experiments, each with duplicated controls. (C) LMB dose dependently reduces anisosome number. Anisosome were induced in TDP-43 2KQ-Clover cells followed by treatment with LMB at the indicated concentrations for 16 h. The histogram shows the distribution of cells (50-90 cells/condition) by anisosome count. (D) FRAP experiments demonstrate that anisosomes remain in a liquid phase following LMB treatment (200 nM, 16 h). Anisosome-induced (24 h) cells were treated with DMSO or LMB for 16 h and then photobleached at the indicated areas. Scale bar 1 µm. (E) Quantification of the FRAP experiment in D. N=18 for control and 16 anisosomes for LMB-treated cells. (F) Time-lapse confocal microscopy detects anisosome fusion after LMB (200 nM, 5 h) treatment in TDP-43 2KQ-Clover cells. Arrows indicate fusion events. Scale bar, 1 µm (G) Overexpression of mCherry-tagged XPO-1 in TDP-43 2KQ-Clover cells induces cytoplasmic TDP-43 puncta. TDP-43 2KQ-Clover cells (green) transfected with mCherry-XPO1 (red) were stained with DAPI (blue) to label nuclei (dashed lines). Cells were imaged 48 h post-transfection. Panels 1, 2 show a representative confocal section, while panels 3-6 show reconstructed 3-D views. The position and volume of anisosomes were also presented in magenta in a surface-rendered view (SRV) in panel 5. The arrow in panel 3 indicates an example of cytoplasmic TDP-43 aggregate. Scale bar, 10 µm. (H) Quantification of the percentage of cells showing cytosolic TDP-43 puncta in XPO-1 positive (pos) and negative (neg) cells in 8 randomly selected fields each with at least 20 cells from 3 independent experiments. **** p <0.0001 by two-tailed unpaired Student’s t-test. (I) FRAP-based confocal imaging reveals the transition of anisosomes into a gel-like state upon XPO-1 overexpression. Scale bar, 1 µm. (J) Anisosome formation changes endogenous XPO-1 localization. TDP-43 2KQ-Clover before or after anisosome induction were stained with anti-XPO-1 antibodies (red) and DAPI (blue). The bleached areas were indicated by dashed lines. Scale bar, 10 µm. (K) Quantification of nuclear endogenous XPO-1 mean fluorescence intensity (MFI) in individual cells (indicated by dots) before or after anisosome induction. **** p < 0.001 by two-tailed unpaired Student’s t-test. N = 58 cells for uninduced and 63 cells for induced condition in 2 independent repeats. (L) A schematic model depicting how the nuclear XPO-1 activity influences TDP-43 liquid-to-phase transition. AS, anisosome.

Earlier version: To corroborate the inhibitor study, we overexpressed mCherry-tagged XPO1. Surprisingly, we observed fewer and enlarged TDP-43 positive puncta in mCherry-XPO1-expressing cells (Figure 5G, panels 1, 2). At first glance, this phenotype appeared similar to that of XPO1-inhibited cells. However, careful examination revealed several distinct features. First, despite the size increase, we did not observe the typical hollowed anisosome ring structure. Instead, these large puncta appeared more irregular in shape. Secondly, while anisosomes in mCherry-XPO1 negative cells were almost exclusively nuclear, ∼30% of mCherry-XPO1-expressing cells contained cytoplasm-localized TDP-43 puncta (Figure 5G, panels 1, 2, 3, Figure 5H). Thirdly, FRAP experiments demonstrated that Clover-TDP-43 fluorescence did not recover after photobleaching, suggesting that TDP-43 was in a gel-like state in transition to aggregates (Figure 5I). Notably, mCherry-XPO1 could be detected within TDP-43 puncta in some cells (Figure 5G, panels 3-6). These findings are consistent with a previous report that showed the regulation of TDP-43 nuclear egress by multiple nuclear export receptors including XPO1 31. As expected, over-expression of mCherry did not alter the subcellular localization of TDP-43 puncta, nor did mCherry co-localize with TDP-43 (Supplemental Figure S4).

Revised version: To corroborate the inhibitor study, we overexpressed mCherry-tagged XPO1, a condition known to enhance nuclear export of XPO1 cargos (Antonarakis, 2018; Kim et al., 2022; Tang et al., 2025). Interestingly, we observed enlarged TDP-43 positive puncta in mCherry-XPO1-expressing cells with characteristics distinct from that of XPO1-inhibited cells (Figure 5G, panels 1, 2). First, we did not observe the typical hollowed anisosome ring structures in mCherry-XPO1-expressing cells. Instead, these large puncta appeared to have irregular shapes. Secondly, while anisosomes in mCherry-XPO1 negative cells were almost exclusively nuclear, ∼30% of mCherry-XPO1-expressing cells contained cytoplasm-localized TDP-43 puncta (Figure 5G, panels 1, 2, 3, Figure 5H). Thirdly, FRAP experiments demonstrated that Clover-TDP-43 fluorescence in mCherry-XPO1-expressing cells did not recover after photobleaching, suggesting that TDP-43 was in a gel-like state in transition to aggregates (Figure 5I). As expected, over-expression of mCherry did not alter the subcellular localization of TDP-43 puncta (Supplemental Figure S4). Thus, higher XPO1 activity is associated with increased accumulation of TDP-43 in a gel-like form in the cytoplasm.

Added in the revised version: Immunostaining of XPO1 showed that anisosome induction depleted endogenous XPO1 from the nucleoplasm (Figure 5J, K). In contrast, no XPO-1 accumulated within anisosomes, probably because antibodies could not stain proteins inside anisosomes due to poor accessibility (Yu et al., 2021).

Added in the revised version: Consistent with this interpretation, we noticed the presence of over-expressed mCherry-XPO1 in TDP-43 puncta in some cells (Figure 5G, panels 3-6). Since it is known that XPO1 cannot bind TDP-43 directly (Pinarbasi et al., 2018), our data suggested a possible indirect interplay between TDP-43 and XPO1.

Added in the revised version: However, the precise mechanism requires future elucidation. Collectively, our results suggest that the stability and dynamics of anisosomes are modulated by XPO1; Low XPO1 activity stabilizes TDP-43 in a liquid state and favors larger nuclear condensates, while high XPO1 activity is associated with a transition to gel-like structures (Figure 5L).

Removed in the revised version: Immunostaining of endogenous XPO1 in anisosome-induced cells showed that anisosome induction depleted XPO1 from the nucleoplasm (Figure 5J, K), further suggesting a possible connection between XPO-1 and TDP-43. Since XPO1 does not bind TDP-43 directly 32, an indirect interaction between these molecules might cause the depletion of XPO-1 from nucleoplasm in anisosome-induced cells. However, since antibody staining could not conclusively demonstrate the sequestration of endogenous XPO-1 in anisosomes due to an antibody penetration barrier 22, alternative interpretation cannot be excluded. Collectively, our results suggest that the stability and dynamics of anisosomes are modulated by XPO1-mediated nuclear export: reduced XPO1 activity stabilizes a liquid state and favors larger TDP-43 condensates, whereas increasing XPO1 activity promotes its transition to a gel-like structure (Figure 5J).

Nuclear export inhibition stabilizes TDP-43 anisosomes in an RNA dependent manner

Earlier version: To further dissect the role of nuclear export in TDP-43 phase regulation, we developed a semi-permeabilized cell–based in vitro assay (Figure 6A). To this end, anisosomes were first induced in Clover-tagged TDP-43 2KQ cells, followed by selective permeabilization of the plasma membrane using the pore-forming toxin streptolysin O (SLO). Permeabilization was monitored at 37 °C by confocal microscopy using the membrane-impermeable dye NucSpot, which stained nuclei only upon plasma membrane disruption.

Revised version: To dissect the factor(s) mediating the effect of nuclear export on TDP-43 phase regulation, we developed a semi-permeabilized cell–based in vitro assay (Figure 6A). To this end, anisosomes were first induced in Clover-tagged TDP-43 2KQ cells, followed by selective permeabilization of the plasma membrane using the pore-forming toxin streptolysin O (SLO). Permeabilization was monitored at 37 °C by confocal microscopy using the membrane-impermeable dye NucSpot, which stained nuclei only upon plasma membrane rupture. Notably, cell permeabilization led to a marked reduction in TDP-43–positive nuclear puncta and the overall fluorescence intensity (Figure 6B). Quantitative analysis revealed an inverse correlation between the NucSpot signal and TDP-43 fluorescence intensity (Figure 6C).

The inhibition of XPO-1 stabilizes anisosomes in an RNA-dependent manner.

Earlier version: (A) A schematic diagram of the semi-permeabilized cell assay. SLO, streptolysin O. (B) TDP-43 in anisosomes becomes soluble after cell permeabilization (indicated by NucSpot positive staining in magenta). Note: The intensity of NucSpot dye indicates permeabilization time. Dashed circles highlight cells permeabilized early during the incubation. The cells have fewer anisosomes. Scale bars, 5μm. (C) TDP-43 mean fluorescence intensity (MPI) displays an inverse correlation with the NucSpot dye intensity in B. AU, arbitrary unit. (D) Small TDP-43 positive puncta were reformed after permeabilized DLD1 cells were incubated with cow liver cytosol in the presence of an ATP regeneration system and GTP (100 μM) (ATP/GTP). Cells were fixed and stained with Hoechst (blue) to label nuclei. Scale bars, 5 μm. (E) Quantification of the TDP-43 Mean Fluorescence Intensity (MPI) in cells as shown in D. Error bars, SD. ****, p<0.0001; ***, p<0.001 by unpaired Student’s t-test. N=at least 35 randomly selected cells representing two independent experiments. (F) LMB treatment (200 nM, 16 h) stabilizes anisosomes. Dashed circles indicate permeabilized cells. Scale bars, 5 μm. (G) RNAase T1 treatment destabilizes anisosomes in DKD1 cells pre-treated with LMB (200 nM, 16 h). DLD1 cells induced for TDP-43 expression were treated with LMB (200nM, 16 h). Cells were permeabilized in the absence or presence of RNAase T1 for 40 min before imaging. Images shown in this figure are maximum intensity projection views of confocal z-sections. Scale bars, 5 μm.

Revised version: (A) A schematic diagram of the semi-permeabilized cell assay. SLO, streptolysin O. (B) TDP-43 in anisosomes becomes soluble after cell permeabilization (indicated by NucSpot positive staining in magenta). Note: The intensity of NucSpot dye correlates with the permeabilization time. Dashed circles highlight cells permeabilized early during the incubation. The cells have fewer anisosomes. Scale bars, 5 μm. (C) TDP-43 mean fluorescence intensity (MFI) displays an inverse correlation with the NucSpot dye intensity in B. AU, arbitrary unit. N =52 cells. (D) Small TDP-43 positive puncta were reformed after the incubation of permeabilized DLD1 cells with cow liver cytosol in the presence of ARS and GTP (100 μM) (ATP/GTP) at 37 ℃for 40 min. Cells incubated with cytosol without ATP/GTP or with PBS plus ATP/GTP serve as controls. Cells were fixed and stained with Hoechst (blue) to label nuclei. Scale bars, 5 μm. (E) Quantification of the TDP-43 mean fluorescence intensity (MFI) in cells as shown in D. Error bars, SD. ****, p<0.0001; ***, p<0.001 by unpaired Student’s t-test. N=at least 35 randomly selected cells representing two independent experiments. (F) LMB treatment (200 nM, 16 h) stabilizes anisosomes. Dashed circles indicate permeabilized cells. Scale bars, 5 μm. (G) RNAase T1 treatment destabilizes anisosomes in DKD1 cells pre-treated with LMB (200 nM, 16 h). DLD1 cells induced for TDP-43 expression were treated with LMB (200 nM, 16 h). Cells were permeabilized in the absence or presence of RNAase T1 for 40 min before imaging. Images shown in this figure are maximum intensity projection views of confocal z-sections. Scale bars, 5 μm.

Added in the revised version: The reduction in TDP-43 signal probably resulted from a shift of TDP-43 from a phase-separated high fluorescent state into a soluble state with reduced fluorescence intensity (Zhang et al., 2026). We attributed this phenotype to the depletion of cytosolic factors and ATP during cell permeabilization because it is known that anisosome formation and maintenance require HSP70, a cytosolic ATPase (Yu et al., 2021). To test this idea, we incubated permeabilized cells with purified cytosol together with an ATP-regenerating system (ARS) and GTP to reconstitute a condition resembling intact cells. Indeed, we observed reformation of TDP-43 puncta after incubation at 37 ℃ (Figure 6D). Although the puncta were smaller than anisosomes, their formation significantly increased TDP-43 fluorescence intensity (Figure 6E). Addition of buffer with ARS and GTP modestly increased TDP-43 puncta and fluorescence intensity, while cytosol alone had no effect (Figure 6D, E). Although the size limitation prevented further characterization of these puncta, the fact that their formation depends on energy and cytosolic factors is consistent with the requirement of HSP70 in anisosome formation (Yu et al., 2021).

Removed in the revised version: Notably, permeabilization led to a marked reduction in TDP-43–positive nuclear puncta and overall fluorescence intensity (Figure 6B). Quantitative analysis revealed an inverse correlation between NucSpot signal and TDP-43 fluorescence (Figure 6C). The reduction in TDP-43 signal was not due to protein elimination. Instead, TDP-43 appeared to shift from a phase-separated state to a soluble state, likely because of cytosolic factor loss and ATP depletion during cell permeabilization. Consistent with this interpretation, supplementation of permeabilized cells with cow liver cytosol together with an ATP-regenerating system (ARS) and GTP restored TDP-43 puncta (Figure 6D), albeit much smaller than anisosomes. Nevertheless, this led to a significantly increase in fluorescence intensity (Figure 6E).

Removed in the revised version: The re-formed puncta may represent intermediates en route to anisosome assembly. However, these structures failed to fuse into larger anisosomes even after prolonged incubation, suggesting that cytosol add-back only partially restores the intracellular environment required for full anisosome maturation. Addition of buffer supplemented with ARS and GTP also modestly induced the formation of small TDP-43 puncta and increased fluorescence intensity (Figure 6D, E). In contrast, buffer or cytosol alone was insufficient to support puncta formation. Together, these results establish this semi-permeabilized system as a platform in which TDP-43 condensates can be dissolved and partially reconstituted under controlled conditions.

Earlier version: We next applied this assay to examine the mechanism underlying leptomycin B (LMB)-induced enlargement of anisosomes. Strikingly, anisosomes in LMB-treated cells remained stable following permeabilization (Figure 6F). However, inclusion of RNase T1, which degrades single-stranded RNA, during permeabilization led to their dissolution (Figure 6G). These findings suggest that LMB stabilizes anisosomes, at least in part, by increasing nuclear RNA availability.

Revised version: We next applied this assay to examine the mechanism underlying LMB-induced enlargement of anisosomes. Strikingly, anisosomes in LMB-treated cells remained stable following permeabilization (Figure 6F). However, inclusion of RNase T1, which degrades single-stranded RNA during permeabilization led to their dissolution (Figure 6G). These findings suggest that LMB stabilizes anisosomes at least in part by increasing nuclear RNA availability.

Nuclear export inhibition mitigates TDP-43 hyperphosphorylation in an organoid model of TDP-43 proteinopathy

Earlier version: Since increasing the XPO1 activity causes more cytoplasmic, gel-like TDP-43, we postulate that reducing XPO1 activity might mitigate TDP-43 cytosolic aggregation. To test this idea, we used a recently established anisosome-bearing organoid model of TDP-43 proteinopathies 16, which is derived from an induced pluripotent stem cell (iPSC) line carrying the ALS-associated mutation (K181E) in the endogenous TDP-43 locus 16. Like the 2KQ mutant, the K181E substitution disrupts the RNA binding activity of TDP-43 33. Additionally, the K181E mutant was reported to accumulate in a liquid-condensed phase in association with HSP70 in cultured cells under overexpression conditions, and like anisosomes 22, these condensates can be converted to solid aggregates during cellular stress or HSP70 depletion 34.

Revised version: Since XPO1 inhibition causes TDP-43 to accumulate in a demixed liquid state, we asked whether reducing XPO1 activity could mitigate TDP-43 cytosolic phosphorylation. To test this idea, we used a recently established organoid model of TDP-43 proteinopathies (Zhang et al., 2025b), which is derived from an induced pluripotent stem cell (iPSC) line carrying the ALS-associated mutation (K181E) in the endogenous TDP-43 locus (Zhang et al., 2025b). Like the 2KQ mutant, the K181E substitution disrupts the RNA binding activity of TDP-43 (Chen et al., 2019). Additionally, the K181E mutant was reported to accumulate in a condensed liquid phase in association with HSP70 in cultured cells under overexpression conditions, and like anisosomes, these condensates can be converted to solid aggregates during cellular stress or HSP70 depletion (Gu et al., 2021).

Earlier version: We treated organoids bearing the endogenous K181E mutation or wild-type (WT) controls with KPT-276 at a low dose (20 nM) for 35 days starting on day 87 based on pilot data indicating good tolerance with prolonged treatment. We then sectioned and stained these organoids with antibodies against Serine 409/410-phosphorylated TDP-43 (p-TDP-43) and total TDP-43 since hyperphosphorylated TDP-43 (S409/410) is a clinically validated pathological marker for TDP-43-containing cytoplasmic aggregates. In this experiment, homozygous organoids were chosen to enhance the detection of endogenous p-TDP-43. Confocal imaging revealed that WT and K181E organoids had similar amounts of total TDP-43, but p-TDP-43 was only significantly present in mutant organoids (Figure 7A), as reported 16. KPT-276 did not alter the total TDP-43 levels in either genotype (Figure 7A) but significantly reduced the p-TDP-43 positive puncta in K181E organoids (Figure 7A, B). This result suggests that nuclear export is required for the formation of p-TDP-43-containing aggregates in a disease relevant organoid model.

Revised version: We treated organoids bearing the endogenous K181E mutation or wild-type (WT) controls with KPT-276 at a low dose (20 nM) for 35 days starting on day 87 based on pilot data indicating good tolerance with prolonged treatment. Homozygous mutant organoids were used to facilitate phosphor-TDP-43 detection. We sectioned and stained these organoids with antibodies against Serine 409/410-phosphorylated TDP-43 (p-TDP-43) and total TDP-43 since hyperphosphorylated TDP-43 (S409/410) is a clinically validated pathological marker for TDP-43-containing cytoplasmic aggregates. Confocal imaging detected significant p-TDP-43 signal in mutant but not in WT organoids (Figure 7A, B), as reported previously (Zhang et al., 2025b). By contrast, we did not detect an increase in total TDP-43 staining (green) in mutant organoids (Figure 7C). Interestingly, KPT-276 treatment significantly decreased p-TDP-43 signal in K181E organoids without affecting the total TDP-43 signal (Figure 7A, B, C). As expected, mutant organoids had more cells with reduced nuclear TDP-43 signal compared to WT organoids (Figure 7A, right panels, D). KPT-276 treatment significantly rescued this phenotype (Figure 7D). These results imply that maintaining TDP-43 in a nuclear demixed liquid state mitigates p-TDP-43 accumulation but whether the reduction of pTDP-43 immunoreactivity reflects a decrease in insoluble TDP-43 aggregates remains to be determined.

Inhibition of XPO-1 mitigates TDP-43 hyperphosphorylation in TDP-43 K181E organoids.

Earlier version: (A) Confocal fluorescence imaging reveals significant reduction in phosphorylated TDP-43 (p-TDP43) in K181E/K181E organoids following KPT-276 treatment. Organoids of the indicated genotypes at day 87 were treated with DMSO as a control, or with KPT-276 (20 nM) for 35 days. Organoids were fixed, sectioned, and stained with antibodies against TUJ1, a neuronal marker (Magenta), TDP-43 (green) and p-TDP-43 (red). n= 3-5 organoids from two individual batches. (B) Quantification of p-TDP-43 Mean Spot Volume (top) and percentage of cells lacking nuclear TDP-43 (bottom) in experiments represented by A. Error bars indicate mean ±s.e.m. * p < 0.05, **** p < 0.0001 by one-way ANOVA, n=3 organoids/each.

Revised version: (A) Confocal fluorescence imaging reveals significant reduction in phosphorylated TDP-43 (p-TDP43) in K181E/K181E organoids following KPT-276 treatment. Organoids of the indicated genotypes at day 87 were treated with DMSO as a control, or with KPT-276 (20 nM) for 35 days. Organoids were fixed, sectioned, and stained with antibodies against TUJ1, a neuronal marker (Magenta), TDP-43 (green) and p-TDP-43 (red). n= 3-5 organoids from two individual batches. (B, C, D) Quantification of p-TDP-43 puncta mean volume (B), total TDP-43 intensity (C) and the percentage of cells with reduced nuclear TDP-43 (D) from experiments represented by A. For B, p-TDP-43-positive puncta were segmented from collected 3-D images for mean puncta volume calculation. Error bars indicate mean ± SEM. * p < 0.05, **** p < 0.0001, ns, not significant, by one-way ANOVA. n = 3 organoids per condition.

Discussion

Earlier version: ALS- and FTD-associated TDP-43 mutants defective in RNA binding are known to undergo self-association through their intrinsically disordered regions and RNA-binding domains, generating demixed liquid, gel, and aggregated states that remain in dynamic exchange. This protein phase behavior is shared by many RNA binding proteins and is increasingly recognized as a contributor to protein aggregation in neurodegenerative diseases 9. Intriguingly, TDP-43 phase separation forms unique layered nuclear structures, with an HSP70-filled central core surrounded by a protein shell made of TDP-43. Although it is unclear how these structures are nucleated in the cell, our observations suggest that, once formed, they can grow by recruiting additional TDP-43 from the surrounding environment. Small anisosomes may also fuse with one another, although such fusion events become infrequent and difficult to detect after anisosomes reach a steady state.

Revised version: ALS- and FTD-associated TDP-43 mutants defective in RNA binding are known to undergo self-association through their intrinsically disordered regions and RNA-binding domains, generating demixed liquid, gel, and aggregated states in dynamic exchange. This protein phase behavior is shared by many RNA binding proteins and is increasingly recognized as a contributor to protein aggregation in neurodegenerative diseases (Shorter, 2019). Intriguingly, TDP-43 phase separation can form unique layered nuclear structures named anisosome. Although it is unclear how these structures are nucleated in the cell, our study suggests that, once formed, anisosomes can grow by recruiting additional TDP-43 from the surrounding environment. Small anisosomes may also fuse with each other, although such fusion events become less frequent after anisosomes mature.

Earlier version: Our study supports the notion that TDP-43, at high concentrations, can change from a demixing-prone, liquid state into a gel-like state en route to forming amorphous aggregates 12, 35. We found that proteasome and HSP90 inhibition or enhancing XPO-1-mediated nuclear egression promotes the conversion of TDP-43 anisosomes into a gel-like state. By contrast, inhibition of RNA splicing, protein translation, or nuclear export all favors the liquid state of TDP-43.

Revised version: Our study supports the notion that TDP-43, at high concentrations, can change from a demixed liquid state into a gel-like state en route to amorphous aggregation (Babinchak et al., 2019; Carey and Guo, 2022). We found that proteasome and HSP90 inhibition or enhancing XPO-1-mediated nuclear egression promotes the conversion of TDP-43 anisosomes into a gel-like state. By contrast, inhibition of RNA splicing, protein translation, or nuclear export all favors a liquid state.

Removed in the revised version: The roles of RNA splicing and nuclear export in TDP-43 phase regulation might be mechanistically linked because anisosomes stabilized upon nuclear export inhibition are sensitive to RNAase treatment in semi-permeabilized cells. Anisosome maturation probably involves one or more RNA splicing intermediate. Therefore, the completion of RNA splicing followed by nuclear egression of spliced RNAs would disfavor anisosome growth and maturation.

Removed in the revised version: Because the TDP-43 2KQ mutant lacks RNA binding activity, anisosome maturation is likely modulated by other RNA-binding factors. In this regard, TDP-43 is known to interact with many splicing factors including members of the hnRNP family like hnRNP A1/A2 and PSF 36; Inhibition of RNA splicing might alter TDP-43’s interactions with these factors, trapping it in an RNA-containing complex that favors anisosome growth and maturation, This could explain why both PlaB and LMB reduce the number of TDP-43 anisosomes while maintaining their dynamic liquid property. Future studies are required to elucidate the precise mechanism by which nuclear transport and RNA splicing modulate anisosome dynamics.

Earlier version: Our validation of proteasome inhibition as a driver of TDP-43 droplet enlargement and hardening reinforces the idea that impaired protein clearance promotes the transition from liquid condensates to gel-like or solid aggregates. This is consistent with previous studies in ALS models, which showed that proteasome stress accelerates TDP-43 aggregation and toxicity 25. In contrast, inhibition of deubiquitinases reduced the number of TDP-43 anisosomes, probably by promoting the proteasome-mediated clearance of TDP-43. The role of HSP90 in this process may be linked to TDP-43 folding stability, as HSP90 can bind TDP-43 to prevent promiscuous interactions 37. Collectively, these observations suggest that TDP-43 condensates are maintained by a delicate balance between formation and clearance, modulated by proteolytic and chaperone activities.

Revised version: Our validation of proteasome inhibition as a driver of TDP-43 droplet enlargement and hardening reinforces the idea that impaired protein clearance promotes the transition of TDP-43 from liquid condensates to gel-like or solid aggregates. This is consistent with previous studies in ALS models, which showed that proteasome stress accelerates TDP-43 aggregation and toxicity (van Eersel et al., 2011). In contrast, inhibition of deubiquitinases reduced the number of TDP-43 anisosomes, probably by promoting the proteasome-mediated clearance of TDP-43. The role of HSP90 in this process may be linked to TDP-43 folding or stability, as HSP90 can bind TDP-43 to prevent promiscuous interactions (Lin et al., 2021). Collectively, these observations suggest that TDP-43 condensates are maintained by a delicate balance between formation and clearance, modulated by proteolytic and chaperone activities.

Added in the revised version: While our chemical and genetic screens have identified multiple pathways influencing TDP-43 phase states, a particularly intriguing aspect of our study is the discovery that nuclear transport modulates the TDP-43 liquid-to-solid transition. However, given the lack of detectable interaction between XPO-1 and TDP-43 as suggested previously (Duan et al., 2022; Pinarbasi et al., 2018), XPO1 probably modulates TDP-43 phase behavior via an indirect mechanism. Indeed, our experiments using semi-permeabilized cells suggest that nuclear transport might affect TDP-43 phase behavior by altering the abundance of nuclear RNA. Anisosome maturation might involve one or more mRNA splicing intermediates. If so, the completion of mRNA splicing followed by nuclear egression of spliced RNAs would disfavor anisosome growth and maturation. However, our data does not exclude other RNA species or RNA-binding proteins as anisosome stabilizer. Whether RNA-dependent stabilization of anisosomes operates in the same way in intact cells also requires further validation. Nevertheless, our findings hint at a potential feedback mechanism in which TDP-43 demixing perturbs nuclear export, analogous to how cytoplasmic TDP-43 aggregates disrupt nuclear import.

Added in the revised version: Because the TDP-43 2KQ mutant lacks RNA binding activity, anisosome maturation is likely modulated by other RNA-binding factors. In this regard, TDP-43 is known to interact with many splicing factors including members of the hnRNP family like hnRNP A1/A2 and PSF (Freibaum et al., 2010).

Added in the revised version: RNA splicing inhibition might alter TDP-43’s interactions with these factors, trapping it in an RNA-containing complex that favors anisosome growth and maturation, This model could explain why both PlaB and LMB reduce the number of TDP-43 anisosomes while maintaining their dynamic liquid property.

Removed in the revised version: While out chemical and genetic screens have identified multiple pathways influencing TDP-43 phase states, a particularly compelling aspect of our study is the discovery that the nuclear export receptor XPO1 governs TDP-43 liquid-to-solid transitions and subcellular localization. Pharmacological inhibition of XPO1 mitigates cytoplasmic p-TDP-43 accumulation, whereas XPO1 overexpression promotes the cytoplasmic accumulation of TDP-43-containing puncta. However, XPO1 may modulate TDP-43 nuclear egression indirectly, as suggested previously 32, 38. While previous studies have not detected strong interactions between XPO-1 and TDP-43 32, a weak or indirect interaction between XPO-1 and TDP-43 may still exist, especially when TDP-43 undergoes demixing. This model would explain why anisosome induction depletes endogenous XPO-1 from the nucleoplasm. Our findings also hint at a potential feedback mechanism in which TDP-43 demixing perturbs nuclear export, analogous to how cytoplasmic TDP-43 aggregates disrupt nuclear import. Given that ALS is characterized by cytoplasmic TDP-43 mis-localization and hyperphosphorylation, the XPO-1-dependent nuclear export pathway may constitute a key molecular switch linking physiological phase behavior to pathological aggregation.

Earlier version: Together, our results support a model in which multiple cellular pathways including proteostasis, RNA metabolism, protein translation, and nuclear export—cooperatively determine the phase behavior of TDP-43. These findings have deepened our understanding of how physiological phase separation can evolve into pathogenic aggregation through cumulative failures of cellular quality control and nuclear export systems.

Revised version: Together, our results support a model in which multiple cellular pathways such as proteostasis network, RNA metabolism, protein translation, and nuclear export cooperatively determine the phase behavior of TDP-43. These findings have deepened our understanding of how physiological phase separation can evolve into pathogenic aggregation through cumulative failures of cellular quality control and nuclear transport systems.

Methods

RNAi-based genetic screen for anisosome regulators

Genome-wide RNA interference (RNAi) screen was performed using DLD1 cells stably expressing Clover-tagged TDP-43 2KQ. Cells were transfected with siRNAs targeting 21,404 genes for 72 hours, followed by induction of anisosome formation with doxycycline (0.5 µg/mL) for 24 h. Specifically, the screen was carried out in 384-well format with 3 individual siRNAs for each gene. To prepare siRNA transfection, control siRNA (40 nmol) was mixed with 2 mL water to create a 50X (20 µM) stock and siRNA-death (20 nmol) was mixed with 1 mL of water for the same concentration. Dilute siRNA stocks to a 400 nM working concentration in water. For each well, mix 0.1 µL transfection reagent in 20 µL of cell culture medium without FBS or P/S. Add 2 µL of siRNA solution to each well and incubate with RNAiMAX for 30 minutes at room temperature. While incubating RNAiMAX, prepare trypsinized cells by ensuring thorough separation and accurate cell counting. Add 20 µL of cell suspension at a density of 0.43 × 105 cells/mL (850 cells/well) in medium containing 20% FBS and 2 x P/S to achieve a final volume of 40 µL per well. Incubate cells for 3 days before addition of 10 µL 10 x doxycycline solution to each well (final volume 50 µL) for anisosome induction. 24 h later, high-throughput (HT) imaging and automated phenotypic screening were conducted to identify genes that modulated anisosome formation. To control for siRNA transfection efficiency, we routinely included a positive knockdown control in which cells were treated with a mixture of siRNAs targeting several essential genes (AllStars Hs Cell Death siRNA (Qiagen, #1027299).

Earlier version: Fluorescence imaging was conducted using an Opera Phenix High-Content Screening System (PerkinElmer). Whole-well images were acquired in confocal mode using a 20X water objective lens after anisosome induction for 24 h. Images were reconstructed and analyzed using the PerkinElmer Columbus server (v2.9.1). Nuclei were segmented based on Hoechst33342 staining, and anisosomes were detected via YFP signals within the nuclear region. The mean anisosome count per nucleus was calculated to assess changes in anisosome levels, while toxicity was evaluated by counting the total number of nuclei. Data were normalized to the siRNA-Neg control for each plate, and Z-scores were calculated to identify significant hits. Identified 1,533 genes whose knockdown significantly reduced the number of anisosomes per cell (Z-score > 2) were subjected to protein network analysis by STRING to identify candidate anisosome regulatory genes that were subject to a second-round screen.

Revised version: Fluorescence imaging was conducted using an Opera Phenix High-Content Screening System (PerkinElmer). Whole-well images were acquired in confocal mode using a 20X water objective lens after anisosome induction for 24 h. Images were reconstructed and analyzed using the PerkinElmer Columbus server (v2.9.1). Nuclei were segmented based on Hoechst33342 staining, and anisosomes were detected via YFP signals within the nuclear region. The mean anisosome count per nucleus was calculated to assess changes in anisosome levels, while toxicity was evaluated by counting the total number of nuclei. Data were normalized to the siRNA-Neg control for each plate. To identify significant hits, Z-scores were calculated by subtracting each anisosome count from the plate sample mean. The resulting values were divided by the plate’s S.D. value. Identified 1,533 genes whose knockdown significantly reduced the number of anisosomes per cell (Z-score > 2) were subjected to protein network analysis by STRING to identify candidate anisosome regulatory genes that were subject to a second-round screen.

A chemical genetic screen for anisosome regulators

We also performed a chemical genetic screen using LOPACR1280 (Sigma) compound library on DLD1 cells expressing Clover-tagged TDP-43-2KQ. Briefly, cells were induced with doxycycline for 24 h to promote anisosome formation, followed by treatment with individual compounds for an additional 24 hours. High-throughput (HT) imaging and automated phenotypic screening were conducted to identify compounds that modulated anisosome formation.

Visualization and analysis of anisosome dynamics

Earlier version: To examine the impact of drug treatment on anisosome size, cells were induced to form anisosome for 24 h and then treated by different inhibitors as follows: cycloheximide (CHX, 20 µg/mL), anisomycin (ANS, 200 nM) for 16 hours. For splicing inhibition, cells were exposed to Pladienolide-B (PlaB) at concentrations of 5 or 20 nM, or DMSO (control) for 16 hours before imaging. Cells were fixed with 4% paraformaldehyde in Phosphate Buffer Saline (PBS). Randomly selected fields were used to quantify the number of anisosomes per cell with at least two biological repeats.

Revised version: To examine the impact of drug treatment on anisosome size, cells were induced to form anisosome for 24 h and then treated by different inhibitors as follows: cycloheximide (CHX, 20 µg/mL), anisomycin (ANS, 200 nM) for 16 hours. For splicing inhibition, cells were exposed to Pladienolide-B (PlaB) at concentrations of 5 or 20 nM, or DMSO (control) for 16 hours before imaging. Cells were fixed with 4% paraformaldehyde in Phosphate Buffer Saline (PBS). Randomly selected fields were used to quantify the number of anisosomes per cell with at least two independent repeats.

To image anisosome fusion, live cell imaging was performed by treating cells with Leptomycin B (200 nM) or PlaB (20 nM) for 5 hours and then imaged by confocal 3-D sectioning for 30 minutes to capture anisosome fusion events. For Fluorescence Recovery After Photobleaching (FRAP), cells were treated with DMSO, Leptomycin B, Geldanamycin or PlaB as indicated in figure legends. Either part of an anisosome or entire anisosome was photobleached and imaged. TDP-43 fluorescence recovery after photobleaching was quantified by ImageJ and calculated to assess anisosome dynamics. At least 10 anisosomes for each condition were bleached and analyzed.

To evaluate the time course of Leptomycin B (LMB)-induced anisosome changes, DLD1 TDP-43 2KQ-Clover cells were seeded and induced with 0.5 µg/mL doxycycline for 24 hours. Cells were treated with 10 µM LMB and fixed at 0-, 8-, and 16-hours post-treatment. Fixed cells were stained with Hoechst to visualize nuclei. 10-20 confocal 3-D sections were obtained for each randomly selected field. Images were converted to maximum projected view by ImageJ before anisosome counting and intensity measurement. The density plot of anisosome counts per cell was generated using R, illustrating dose-dependent effects of leptomycin on anisosome formation.

To test the effect of XPO1 overexpression on TDP-43 phase regulation, DLD1 TDP-43 2KQ-Clover cells were co-transfected with mCherry-tagged XPO1 (mCh-XPO1) for 24 h before doxycycline was added at 0.5 µg/mL to induce anisosome formation for 24 h before confocal imaging.

To study the relationship between anisosome induction and endogenous XPO1 localization, cells were seeded, induced with 0.5 µM doxycycline for 48 hours, and fixed with paraformaldehyde in PBS. Immunostaining was performed using an XPO1 antibody (Cell Signaling, 46249S, 1:125) and Hoechst nuclear stain. Fluorescence intensity of XPO1 in both induced and uninduced cells was measured by ImageJ.

Semi-permeabilized cell experiments

Earlier version: To reconstitute anisome disassembly and re-assembly in seminar permeabilized cells, 30,000 DLD1 cells were seeded in an eight-well ibidi cell chamber two days before the experiment. Cells were treated with Doxycycline for 24 h to induce anisosomes and then treated with inhibitors as indicated in the figure legends. We then washed Cells three times with ice-cold PBS containing 2 mM MgCl2 (PBS-Mg) and then treated these cells in the same buffer (300 µl) containing 200 units of SLO (Sigma-Aldrich, cat. no. SAE0089) on ice for 30 min. The cells were then washed two times with PBS-Mg buffer and incubated with 200 µl of reagent mixture containing 200 µl of PBS in the absence of presence of ARS/GTP or 200 µl of cow liver cytosol (CLC) in the presence of the NucSpot dye. Cells were either imaged directly or fixed with 4% paraformaldehyde before imaging by a Nikon CSU-W1 SoRa super-resolution confocal microscope.

Revised version: To reconstitute anisome disassembly and re-assembly in seminar permeabilized cells, 30,000 DLD1 cells were seeded in an eight-well ibidi cell chamber two days before the experiment. Cells were treated with Doxycycline for 24 h to induce anisosomes and then treated with inhibitors as indicated in the figure legends. We then washed cells three times with ice-cold PBS containing 2 mM MgCl2 (PBS-Mg) and then treated these cells in the same buffer (300 µl) containing 200 units of SLO (Sigma-Aldrich, cat. no. SAE0089) on ice for 30 min. The cells were then washed two times with PBS-Mg buffer and incubated with 200 µl of PBS buffer containing the NucSpot dye at 37 ℃ for 30-40min to permeabilize the cell and disassemble anisosomes. After that, cells were incubated with 200 µl cow liver cytosol (Lee et al., 2025; Ye et al., 2001) in the absence or presence of an ATP regenerating system (Ye et al., 2001) and GTP (100 μM) at 37 ℃ for 30-40 min. Cells were either imaged directly or fixed with 4% paraformaldehyde before imaging by a Nikon CSU-W1 SoRa spinning disk confocal microscope.

Experiments with organoids

Earlier version: Stem cell gene editing of the patient K181E mutation in iPSC cells was reported previously 16. Differentiation of iPSCs into forebrain neurons was performed following a previously published protocol39. Briefly, the procedure used an induction medium (IM-N2) and a neuronal culture medium (CM). IM-N2 was prepared with Knockout DMEM/F12, N2 supplement, NEAA, Gluta-MAX, Chroman I, and doxycycline. On Day 0, iPSCs were observed for confluency, dissociated with Accutase, centrifuged, and resuspended in IM-N2 with Chroman I, and then seeded into Matrigel-coated plates. Over the next four days, cells were monitored microscopically for neurite extensions while media containing doxycycline is refreshed daily. By Day 4, cells exhibit neurite growth and are ready for replating onto poly-L-ornithine (PLO)-coated dishes, prepared in advance by coating with PLO solution, incubating, washing, and drying. Replated cells were cultured in CM comprising BrainPhys medium, B27+ supplement, neurotrophic factors (GDNF, BDNF, NT-3), laminin, and doxycycline, to promote neuronal maturation.

Revised version: Stem cell gene editing of the patient K181E mutation in iPSC cells was reported previously (Zhang et al., 2025b). Differentiation of iPSCs into forebrain neurons was performed following a previously published protocol (Fernandopulle et al., 2018). Briefly, the procedure used an induction medium (IM-N2) and a neuronal culture medium (CM). IM-N2 was prepared with Knockout DMEM/F12, N2 supplement, NEAA, Gluta-MAX, Chroman I, and doxycycline. On Day 0, iPSCs were observed for confluency, dissociated with Accutase, centrifuged, and resuspended in IM-N2 with Chroman I, and then seeded into Matrigel-coated plates. Over the next four days, cells were monitored microscopically for neurite extensions while media containing doxycycline is refreshed daily. By Day 4, cells exhibit neurite growth and are ready for replating onto poly-L-ornithine (PLO)-coated dishes, prepared in advance by coating with PLO solution, incubating, washing, and drying. Replated cells were cultured in CM comprising BrainPhys medium, B27+ supplement, neurotrophic factors (GDNF, BDNF, NT-3), laminin, and doxycycline, to promote neuronal maturation.

Earlier version: 3-D culture and organoid growth were performed as previously described 40. In brief, iPSCs grown on Matrigel were dissociated into single cell suspension by Versene solution (ThermoFisher) and seeded into a 12-well Aggrewell plate (Stemcell Technologies) at 4,000 cells/ well. Next day, spheroids were transferred to an ultralow attachment plate (Corning) containing phase I medium: DMEM/F12, 20% Gibco KnockOut Serum Replacement, 1X Glutamax, 1X MEM Nonessential Amino Acid, 55 µM β-Mercaptoethanol, 1X Pen/Strep, 2 µM Dorsomorphine, 2 µM A83-01. After 5 days, medium was switched to phase II medium: DMEM/F12, 1X N-2 Supplement, 1X Glutamax, 1X MEM Nonessential Amino Acid solution, 1X Pen/Strep, 4 ng/mL WNT3a, 1 µM CHIR-99021, 1 µM SB-431542. On day 7, spheroids were embedded into Matrigel and allowed to continue growing for 7 more days. On day 14, individual spheroids were manually freed from the Matrigel and transferred to a SpinOmega bioreactor spinning at 120 RMP with phase III media: DMEM/F12, 1X N-2 Supplement, 1X B-27 Supplement, 1X Glutamax, 1X MEM Nonessential Amino Acid, 55 µM 2-Mercaptoethanol, 1X Pen/Strep, 2.5 µg/mL insulin. 50 days later, the medium was switched to final differentiation medium: Neurobasal medium, 1X B-27 Supplement, 1X Glutamax, 55 µM 2-Mercaptoethanol, 1X Pen/Strep, 0.2 mM Ascorbic acid, 0.5 mM cAMP, 20 ng/mL brain-derived neurotrophic factor (BDNF), 20 ng/mL glial-derived neurotrophic factor (GDNF). 107-day-old organoids were dissociated into single cells using a 50:50 mixed of Accutase and 0.25% Trypsin with DNase I (1 mg/mL) and plated onto chambered glass slides pretreated with 1% Matrigel (Corning, Inc). Alternatively, organoids were fixed and sectioned for immunostaining (see below). For drug treatment, forebrain organoids (87 Day) were transferred to a 12-well plate on an orbital shaker (120 RPM) and treated with 20 nM KPT276 in final differentiation medium. Medium was changed every other day (Dexoregen, Inc). The organoids were harvested and frozen after 35 days of treatment for sectioning and immunostaining. Each experiment was performed with two batches of organoids per condition to reduce the effect of batch variation. The wildtype cells and K181E mutant have the same genetic background.

Revised version: 3-D culture and organoid growth were performed as previously described (Nguyen, 2022). In brief, iPSCs grown on Matrigel were dissociated into single cell suspension by Versene solution (ThermoFisher) and seeded into a 12-well Aggrewell plate (Stemcell Technologies) at 4,000 cells/ well. Next day, spheroids were transferred to an ultralow attachment plate (Corning) containing phase I medium: DMEM/F12, 20% Gibco KnockOut Serum Replacement, 1X Glutamax, 1X MEM Nonessential Amino Acid, 55 µM β-Mercaptoethanol, 1X Pen/Strep, 2 µM Dorsomorphine, 2 µM A83-01. After 5 days, medium was switched to phase II medium: DMEM/F12, 1X N-2 Supplement, 1X Glutamax, 1X MEM Nonessential Amino Acid solution, 1X Pen/Strep, 4 ng/mL WNT3a, 1 µM CHIR-99021, 1 µM SB-431542. On day 7, spheroids were embedded into Matrigel and allowed to continue growing for 7 more days. On day 14, individual spheroids were manually freed from the Matrigel and transferred to a SpinOmega bioreactor spinning at 120 RMP with phase III media: DMEM/F12, 1X N-2 Supplement, 1X B-27 Supplement, 1X Glutamax, 1X MEM Nonessential Amino Acid, 55 µM 2-Mercaptoethanol, 1X Pen/Strep, 2.5 µg/mL insulin. 50 days later, the medium was switched to final differentiation medium: Neurobasal medium, 1X B-27 Supplement, 1X Glutamax, 55 µM 2-Mercaptoethanol, 1X Pen/Strep, 0.2 mM Ascorbic acid, 0.5 mM cAMP, 20 ng/mL brain-derived neurotrophic factor (BDNF), 20 ng/mL glial-derived neurotrophic factor (GDNF). 107-day-old organoids were dissociated into single cells using a 50:50 mixed of Accutase and 0.25% Trypsin with DNase I (1 mg/mL) and plated onto chambered glass slides pretreated with 1% Matrigel (Corning, Inc). Alternatively, organoids were fixed and sectioned for immunostaining (see below). For drug treatment, forebrain organoids (87 Day) were transferred to a 12-well plate on an orbital shaker (120 RPM) and treated with 20 nM KPT276 in final differentiation medium. Medium was changed every other day (Dexoregen, Inc). The organoids were harvested and frozen after 35 days of treatment for sectioning and immunostaining. Each experiment was performed with two batches of organoids per condition to reduce the effect of batch variation. The wildtype cells and K181E mutant have the same genetic background.

Tissue preparation and immunostaining

Organoids were fixed in 4% paraformaldehyde in PBS for 1 hour at room temperature, washed with PBS, and immersed in 15% sucrose/PBS solution overnight. Subsequently, organoids were embedded in O.C.T. compound (Sakura) in a plastic mold and frozen down in an ultralow freezer. Embedded organoids were sliced onto glass slides using a Cryostat (Leica). Slides were rinsed with PBS, permeabilized with 0.5% Triton-X/PBS solution for 1 hour at room temperature and blocked using 1% donkey serum in 0.1% Tween-20/PBS solution for 1 hour. Primary antibodies, chicken anti-TUJ1 (Abcam, Ab41489, 1:1,000 dilution), rabbit anti-TDP-43 (10782-2-AP, 1/1000 dilution), mouse anti-phospho-TDP-43 (Cosmobio, CAC-TIP-PTD-M01A, 1/500 dilution), were added to the slides and the slides were incubated in a humidified chamber at 4 °C overnight. After several washes in 0.1% Tween-20/PBS solution, DAPI (Sigma) or secondary antibodies, donkey anti-rabbit, anti-mouse, and anti-chicken (Jackson ImmunoResearch), diluted 1:1,000 in blocking solution, were added to the slides. After 1 hour of incubation at room temperature, the slides were washed and mounted in antifade mounting solution (Fisher scientific).

Image acquisition and processing

Fluorescence confocal z-section images were acquired using a Nikon CSU-W1 SoRa microscope equipped with temperature and CO2 control enclosures and a 60x TIRF lens. Randomly selected fields were imaged each with ∼25 slices. Maximum intensity projection views were reconstructed and analyzed by the open-source Image J software (also named Fiji). 3-D and time lapse visualizations were achieved by the Imaris software. Fluorescence intensity was quantified using Fiji. To automatedly count anisosomes, maximum intensity projection views were split into individual channels by Fiji. A consistent default threshold method was applied to each channel. The particle analysis function was used to automatically identify anisosomes or other fluorescence structures for size and intensity measurement.

Statistical analyses were conducted using Excel (for Student’s t-test) or GraphPad Prism versions 8.0, 9.0, 11. P-values were calculated using Student’s t-test in Excel or one-way ANOVA in GraphPad Prism. Curve fitting, including linear and nonlinear models, as well as IC50 calculations, were also performed using GraphPad Prism versions 8.0, 9.0 and 11.

Statistic and reproducibility

Earlier version: All statistical analyses were conducted with GraphPad Prism v10 or Excel. Statistical methods and the number of cells or brain organoid samples (N) are indicated in figure legends or shown in figures as individual data point. Biological repeats (n) are specified in figure legends. No biological repeat was excluded from the analyses. All experiments were repeated at least twice with individual data point labeled in figures unless specified. For imaging analyses, cells in randomly selected fields were analyzed. The researchers were not blinded. The iPSC cell line was derived from a ADRDs genetic risk-free clone, which was initially obtained from a male. Figures were prepared using ImageJ 1.54f, Imaris 9.9.0, Adobe Photoshop v25.12.1, and Adobe Illustrator 28.7.4.

Revised version: All statistical analyses were conducted with GraphPad Prism v10 or Excel. Statistical methods and the number of cells or brain organoid samples (N) are indicated in figure legends or shown in figures as individual data point. Independent repeats (n) are specified in figure legends. No experiment was excluded from the analyses. All experiments were repeated at least twice with individual data point labeled in figures unless specified. For imaging analyses, cells in randomly selected fields were analyzed. The researchers were not blinded. The iPSC cell line was derived from a ADRDs genetic risk-free clone, which was initially obtained from a male. Figures were prepared using ImageJ 1.54f, Imaris 9.9.0, Adobe Photoshop v25.12.1, and Adobe Illustrator 28.7.4.